<figure>
  <IMG SRC="https://raw.githubusercontent.com/mbakker7/exploratory_computing_with_python/master/tudelft_logo.png" WIDTH=250 ALIGN="right">
</figure>


# Python Notebook #3

## Table of Contents

<ul>
    <li> <a href="#advanced_strings">3.1 Advanced Strings</a>
    <li> <a href="#advanced_functions">3.2 Advanced Functions</a>
    <li> <a href="#working_with_files">3.3 Working with files</a>
    <li> <a href="#debugging">3.4 Debugging</a>
</ul>

```python
# it is common to have the first cell in a Notebook with 
# all imports needed

from math import pi # access to pi without writing 'math.pi' every time
import os # module that we will use later

```

<div id='advanced_strings'></div><br><h2> 3.1 Advanced Strings</h2><br><div style="text-align: justify">Welcome to the third Notebook. In this Notebook we are going to learn some advanced Python. Let's first start with strings. Run the code below and see what it prints out.<b><b>


```python
MyFString = f"This is an F-String"
MyString = "This is a string"
print(MyFString)
print(MyString)
```

<div style="text-align: justify">
Now let's try inserting some data into our <b><code>print()</code></b> function. We'll use the list of integers <b><code>[1,2,3,4]</code></b>. 

```python
Data = [1,2,3,4]

MyFString = f"Data1: {Data[0]}, Data2: {Data[1]}, Data3: {Data[2]}, Data4: {Data[3]}"

print(MyFString)
print("Data1:",Data[0],",Data2:",Data[1],",Data3:",Data[2],",Data4:",Data[3])
```

<div style="text-align: justify">As you can see from the above code, it is much easier to insert variables in a string by using an f-string (formatted string).

<div class="alert alert-block alert-info">
<b>Exercise 3.1.1</b><br><br>
<div style="text-align: justify">
    Use the <code>print()</code> function and an f-string to make a statement about some car using all variables stored in the <code>car_info</code> dictionary.
</div>

```python
# you do not have to change anything in this cell

car_info = {
    'top_speed' : '229', #km/h
    'type' : 'Opel Astra'
}
```

```python
# type your code here
message = ...

###BEGIN SOLUTION TEMPLATE=
message = f"Car {car_info['type']} can drive at max {car_info['top_speed']} km per hour"
###END SOLUTION

print(message)
```

<!-- grade_id: Q3.1.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert car_info['top_speed'] in message and car_info['type'] in message, '3.1.1 - Incorrect answer'
###END HIDDEN TESTS
```

<h3>Formatting numbers</h3><br><div style="text-align: justify">Using f-strings makes formatting numbers really easy. Just add a colon character after a number value and specify how you want to format the number. The following table demonstrates a couple of examples with the number $1$:<br><br>

| Code | Result|
|------|------|
| 1:.2f | 1.00|
| 1:.0f| 1|
| 1:.10f| 1.0000000000 | 
| 1:%| 100.000000%|
| 1:.1%| 100.0% |
| 1:e| 1.000000e+00 |

<div style="text-align: justify"><br>As you can see the default number of decimal places is six. Furthermore, the <b><code>%</code></b> formatting operator assumes that $1$ is equal to $100$%, which is usual when working with fractions, and the formatting operator <b><code>e</code></b> formats using scientific notation.

<div class="alert alert-block alert-info">
<b>Exercise 3.1.2</b><br><br> Print the value of <b>π</b> using $4$ decimal places.
</div>

```python
# write your code here
rounded_pi = ...

###BEGIN SOLUTION TEMPLATE=
rounded_pi = f"{pi:.4f}"
###END SOLUTION
print(rounded_pi)
```

<!-- grade_id: Q3.1.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert rounded_pi == "3.1416", '3.1.2 - Incorrect answer'
###END HIDDEN TESTS
```

<div style="text-align: justify">
Now let's use our newfound knowledge of strings to make a simple progress bar. During other courses, you'll sometimes have to write algorithms that take a long time to run. In this case, it is useful to have a progress bar. Our example of a progress bar makes use of the <b><code>sleep()</code></b> function, from the <b><code>time</code></b> module, to simulate elapsed time.

```python
import time

for i in range(11):
    print(f"Loading: {i*10}%", )
    time.sleep(0.5)  # Wait for 0.5 seconds
```

<div style="text-align: justify">
This works! Though it is not that pretty to look at. It would look nicer to not have it print a new line each time. This is where <b>escape characters</b> come in. These characters can do some special things in strings. Below an example of some escape characters:


<h3>Escape characters</h3>

| Code | Result|
|------|------|
|   \\'  | ' |
| \\\    | \\ |
| \\n | new line| 
| \\r | carriage return  |
| \\t | tab |
| \\b | backspace |

<div style="text-align: justify"><br>We can use some of these characters in our code. Let's use the carriage return character to make our progress bar not print out a new line every time. We can do this by adding <b><code>end="\r"</code></b> into our print function. The <b><code>end</code></b> keyword specifies a string that gets printed at the end. The string we print at the end here is the carriage return character. This carriage resets the print function to the start of the line; thus making the next print function overwrite the current printed line. Try it and see what happens:

```python
print("Will I get overwritten?", end="\r")
print("This is a very important message")
```

Now let's add this to our progress bar... 

```python
import time
for i in range(11):
    print(f"Loading: {i*10}%", end="\r")
    time.sleep(0.5) # Wait for 0.5 seconds
print("Loading complete!")
```

As you can see, it works beautifully!

<div class="alert alert-block alert-danger"><b>Additional study material:</b>
    
* Official Python Documentation - https://docs.python.org/3/tutorial/inputoutput.html
* https://realpython.com/python-f-strings/
* Think Python (2nd ed.) - Section 14 

<div id='advanced_functions'><br>
<div style="text-align: justify"><h2>3.2 Advanced Functions</h2><br>Sometimes you want to use the same code multiple times, so you could embed this code into a function. However, sometimes the code you want to use is so short that putting it into a function feels a bit over the top. This is where <b>lambda functions</b> are useful. <br><br>Lambda functions are functions that can take any number of arguments but can only have one expression in their function body. To demonstrate, see the code below. Here we have two functions that do exactly the same, but one is a lambda function and the other one is a normal function. 

```python
sqrt_lambda = lambda x : x**0.5

def sqrt(x):
    sqrt = x**0.5
    return sqrt

print(f"The square root of 16 is equal to {sqrt_lambda(16):.0f}")
print(f"The square root of 16 is equal to {sqrt(16):.0f}")
```

<div style="text-align: justify">
As you can see, the lambda version is much more concise. It automatically returns the computed value for you as well.

<div class="alert alert-block alert-info">
<b>Exercise 3.2.1</b><br><br><div style="text-align: justify"> Write a lambda function that converts a number from degrees to radians.
</div>

```python
DegToRad = ...

###BEGIN SOLUTION TEMPLATE=
DegToRad = lambda theta : theta * (pi / 180)
###END SOLUTION

Angle = 180 # Degrees
print(f"An angle of {Angle} Degrees is equal to {DegToRad(180):.3f} radians")
```

<!-- grade_id: Q3.2.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert type(DegToRad) == type(lambda x: x) and abs(DegToRad(90) - pi/2) <= 1e-6, '3.2.1 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info">
<b>Exercise 3.2.2</b><br><br><div style="text-align: justify"> Write a lambda function that takes four inputs: $(x_1, y_1, x_2, y_2)$ and computes the <a href="https://en.wikipedia.org/wiki/Euclidean_distance">Euclidian distance</a> between point 1 $(x_1,y_1)$ and point 2 $(x_2,y_2)$.
</div>

```python
Distance = ...

###BEGIN SOLUTION TEMPLATE=
Distance = lambda x1, y1, x2, y2: sqrt((x1 - x2) ** 2 + (y1 - y2) ** 2)
###END SOLUTION

x1, y1 = 1, 1
x2, y2 = 2, 2

print(f"Distance between points ({x1}, {y1}) and ({x2}, {y2}) is {Distance(x1, y1, x2, y2):.3f}")
```

<!-- grade_id: Q3.2.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert abs(Distance(1, 1, 3, 3) - 2 * sqrt(2)) <= 1e-6, '3.2.2 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-danger"><b>Additional study material:</b>
    
* Official Python Documentation - https://docs.python.org/3/reference/expressions.html
* https://realpython.com/python-lambda/

<div id='working_with_files'><br><h2>3.3 Working with files</h2><br><div style="text-align: justify">A lot of the work you'll do in Python will have the following structure:<br><br>1. Read data from a file<br>2. Perform computations on the data<br>3. Visualize the results and/or save the results to a file<br><br>So far, we have only learned about computations. So let's learn a bit about how to manage files. Actually, opening or saving files is usually done with the help of modules which you will learn in more detail in Notebook 4 and 6. What we'll discuss here is how to manage file paths.<br><h3>File paths</h3><br>
To learn how to use files we need to learn how file paths in computers work. If you are tech-savvy and know how file paths work you can skip this part. <br><br>
File paths in computers work like a tree. They start at the <b>root</b> directory, which is often the <b><code>C:</code></b> drive (in Windows). This is the name of the hard drive that stores your Operating System. From the <b><code>C:</code></b> drive you can navigate into other directories. This is done using the <b><code>\</code></b> character, however in other Operating Systems often the <b><code>/</code></b> delimiter is used.<br> If a file is in the folder <b><code>Users</code></b>, which is stored in the <b><code>C:</code></b> directory, the file path would be <b><code>C:\Users</code></b>. These types of file paths are called <b>absolute paths</b>. This file path is valid for most computers that run Windows, but some other Operating Systems may have different folder setups. This is why it is useful to use <b>relative paths</b>. Relative paths do not start from the root directory. Instead, they start from the directory you are currently in. By default, Jupyter Notebooks are stored in <b><code>C:\Users\CurrentUser</code></b> (where <b><code>CurrentUser</code></b> is your Windows username). To move into a directory using a relative path, for example, to the desktop folder, you would just write <b><code>.\Desktop</code></b>. To move back a directory, using a relative path, you would type <b><code>..</code></b>

```python
import os

# List all the entries in your current directory
print(os.listdir())

#or like this
print(os.listdir('./'))

# List all entries if we go back one level
print(os.listdir('../')) # note that we use the / as delimiter, since a \ won't work on Mac

# Now list all contents of the directory resource, which is one level back
os.listdir('../resource')
```

<div style="text-align: justify">You see that, since you are using Vocareum, the script does not access your computer paths. Using relative (or absolute) paths in Vocareum will result in you navigating through Vocareum's folders, which is useful for the purpose of this course as you will be accessing files that we upload here. The theory remains the same once you migrate your Python skills to your own computer, outside of Vocareum.

<div class="alert alert-block alert-warning"><center>
Keep in mind that, in Python, all file paths must be strings!

<h4><code>pathlib</code> and <code>os</code> modules</h4><div style="text-align: justify"><br>These modules are very useful in managing and navigating your file paths. The function <b><code>path.expanduser('~')</code></b>, from the <b><code>os</code></b> module, allows you to find your root directory (or Vocareum's root directory in this case), independent of your Operating System. Try the below cell to see it.

```python
from pathlib import Path
import os

root_path = os.path.expanduser('~')
print(root_path)
```

<div style="text-align: justify">The path shown above is thus the absolute path to your current directory.

This can come in handy when you write a code that needs to create directories in the user's computer to save data files and/or plots. As an example, the code below checks if a directory exists and, if it doesn't, it creates one.

```python
# Print contents of current directory using the root_path
print('Contents of current directory (before):')
print(os.listdir(root_path))

# the os.path.join is used to concatenate two strings to form a path string with the appropriate delimiter

imdir = os.path.join(root_path,'plots') 
print(f'\nimdir = {imdir}')

# the below line will check if a directory named 'plots' exists in your current directory..
# if not, it will create one

Path(imdir).mkdir(parents=True, exist_ok=True)

print('\nContents of current directory (after creating the new directory):')
print(os.listdir(root_path))
```

```python
# run this cell to delete the folder that was just created
try:
    os.rmdir(imdir)
    print(f'Directory {imdir} has been deleted.')
except:
    print('You already deleted the folder. :)')
```

<div style="text-align: justify">Now you are, hopefully, a bit more used to working with file paths. For the next test, we are going to try to open a file. We can use some built-in Python functions to open a <b><i>*.txt</i></b> file and print its contents.

<div class="alert alert-block alert-info"><div style="text-align: justify"><b>(Fixing) Exercise 3.3.1</b><br><br>
    An engineer from Shell wrote a piece of code to find the file with updated logs. The file has a <code>*.txt</code> extension. It is located in the Vocareum environment and can be accessed using a relative or absolute path. The engineer wrote a piece of code to read the file and print its contents but something went wrong.
</div>

```python
FilePath = "./resource/lib/publicdata/UpdatedLogs.txt"# Path of the file + (filename.txt)

###BEGIN SOLUTION TEMPLATE=
FilePath = "./resource/lib/publicdata/updated_logs.txt"
###END SOLUTION


txt_File = open(FilePath, mode="r")                         # Open the file in read (r) mode
File_Contents = txt_File.read()                             # Read the contents of the file
txt_File.close()                                            # Close the file (good practice to do so!)

print(File_Contents)                                        # Print the file
```

<!-- grade_id: Q3.3.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert FilePath == "./resource/lib/publicdata/updated_logs.txt", '3.3.1 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-danger"><b>Additional study material:</b>
    
* Official Python Documentation - https://docs.python.org/3/library/filesys.html
* https://realpython.com/working-with-files-in-python/
* Think Python (2nd ed.) - Section 14 

<div id='debugging'><br><h2>3.4 Debugging</h2><br><div style="text-align: justify">It is very easy (and common) to make mistakes when programming. We call these errors <i>bugs</i>. Finding these <i>bugs</i> in your program and resolving them is what we call <i>debugging</i>.<h3>Errors</h3><br><div style="text-align: justify">
    According to <a href="https://greenteapress.com/wp/think-python-2e/"><i>Think Python</i></a> — <i>Appendix A</i>, there are three different types of errors:<br><h4>1. Syntax errors</h4><br>
<i>"In computer science, the syntax of a computer language is the set of rules that defines the combinations of symbols that are considered to be correctly structured statements or expressions in that language."</i><br><br>Therefore, a syntax error is an error that does not obey the rules of the programming language. For example, parenthesis always comes in pairs... so <b><code>(1+2)</code></b> is OK, but <b><code>1+2)</code></b> is not. Below another example of a syntax error.  As you will see — this error is caught by the interpreter before running the code (hence, the <code>print</code> statements do not result in anything being printed).

```python
# I want to raise 2 to the 3rd power.
# However, I apply the wrong syntax, causing a syntax error:
print('Message before')
2***3
print('Message after')
```

<h4>2. Runtime errors</h4><div style="text-align: justify"><br><i>"The second type of error is a runtime error. This type of error does not appear until after the program has started running. These errors are also called <b>exceptions</b>, as they usually indicate that something exceptional (and bad) has happened."</i><br><br>Below an example of a runtime error:

```python
# A small script to express fractions as decimals
numerators = [1, 7, 5, 12, -1]
denominators = [6, 8, -1, 0, 5]
fractions = []

for i in range(len(numerators)):
    fractions.append(numerators[i] / denominators[i])
    print(f'New fraction was added from {numerators[i]} and {denominators[i]}!\n It is equal to {fractions[i]:.3f}')
    # Error will appear, since you cannot divide by 0
```

<h4>3. Semantic errors</h4><div style="text-align: justify"><br>According to the Oxford Dictionary, 'semantic' is an adjective relating to meaning. Therefore, a 'semantic error' is an error in the meaning of your code. Your code will still run without giving any error back, but it will not result in what you expected (or desired). For that reason, semantic errors are the hardest to identify. Below an example:

```python
# I want to raise 2 to the 3rd power.
# However, I apply the wrong syntax that does not represent "pow()". 
# No error message is created, because this syntax is used 
# for another function in Python
# However, this results in an output I did not expect nor desire

power_of_2 = 2^3
print(f'2 to the 3rd power is {power_of_2}')
```

<h3>Debugging strategies</h3><div style="text-align: justify"><br>There are a few ways to debug a program. A simple one is to debug by tracking your values using print statements. By printing the values of the variables in between, we can find where the program does something unwanted. For example, the code block below:

```python
A = [0, 1, 2, 3]

def sumA(my_list):
    "returns the sum of all the values in a given list"
    my_sum = 0
    i = 0
    while i < len(A):
        my_sum = A[i]
        i += 1
    return my_sum

print('The sum of the elements of the list A is {}.'.format(sumA(A)))
```

<div style="text-align: justify">We see that our <b><code>sumA()</code></b> function outputs $3$, which isn't the sum of the contents of the list $A$. By adding a <b><code>print(my_sum)</code></b> inside the loop we can get a clearer understanding of what goes wrong.

```python
def sumA(my_list):
    "returns the sum of all the values in a given list"
    my_sum = 0
    i = 0
    while i < len(A):
        my_sum = A[i]
        print('var my_sum[{}] = {}'.format(i,my_sum))
        i += 1
    return my_sum

print('The sum of the elements of the list A is {}.'.format(sumA(A)))
```

<div style="text-align: justify">It looks like the function is just stating the values of the list $A$, but not adding them... so we must have forgotten to add something. Below the fixed version of that function.

```python
def sumA_fixed(my_list):
    "returns the sum of all the values in a given list"
    my_sum = 0
    i = 0
    while i < len(A):
        my_sum += A[i]
        print('var my_sum[{}] = {}'.format(i,my_sum))
        i += 1
    return my_sum

print('The sum of the elements of the list A is {}.'.format(sumA_fixed(A)))
```

<div class="alert alert-block alert-info"><b>(Fixing) Exercise 3.4.1</b><br><br><div style="text-align: justify">Fix the syntax errors so it prints "AES" without removing the variable that holds it. You'll need to fix 2 errors.

```python
def get_abbreviation():
    my abbreviation = "AES"
     return my_abbreviation
        
###BEGIN SOLUTION TEMPLATE=
def get_abbreviation():
    my_abbreviation = "AES"
    return my_abbreviation
###END SOLUTION
        
print(get_abbreviation())
```

<!-- grade_id: Q3.4.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert get_abbreviation() == "AES", '3.4.1 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>(Fixing) Exercise 3.4.2</b><br><br><div style="text-align: justify">Solve the runtime error so all values of $B$ are printed. The output after running <code>create_string_from_lists()</code> should be the following:<br>
<code>
A[0] = 2
B[0] = 5
A[1] = 3
B[1] = 6
A[2] = 4
B[2] = 7
B[3] = 8
</code>

```python
def create_string_from_lists():
    s = ""
    
    A = [2, 3, 4]
    B = [5, 6, 7, 8]
    for i in range(4):
        s += f"A[{i}] = {A[i]}\n"
        s += f"B[{i}] = {B[i]}\n"
    print(s)

###BEGIN SOLUTION TEMPLATE=
def create_string_from_lists():
    s = ""
    
    A = [2, 3, 4]
    B = [5, 6, 7, 8]
    for i in range(4):
        
        if i < 3:
            s += f"A[{i}] = {A[i]}\n"
            s += f"B[{i}] = {B[i]}\n"
        else:
            s += f"B[{i}] = {B[i]}\n"
    print(s)   
###END SOLUTION

create_string_from_lists()
```

<!-- grade_id: Q3.4.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert "B[3] = 8" in create_string_from_lists(), '3.4.2 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>(Fixing) Exercise 3.4.3</b><br><br><div style="text-align: justify">Find the semantic error in this function:

```python
def factorial(x):
    "returns the factorial of x"
    if x == 0:
        return 1
    else: 
        return x ** factorial(x-1)

###BEGIN SOLUTION TEMPLATE=
def factorial(x):
    "returns the factorial of x"
    if x == 0:
        return 1
    else: 
        return x * factorial(x-1)
###END SOLUTION

factorial(4)
```

<!-- grade_id: Q3.4.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert factorial(5) == 120, '3.4.3 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-danger"><b>Additional study material:</b>
    
* Official Python Documentation - https://docs.python.org/3/library/debug.html
* Think Python (2nd ed.) - Appendix A and all (sub-)sections

<h4>After this Notebook you should be able to:</h4>

- print a variable, formatting it in an appropriate manner
- know the existence of escape characters
- know how to use lambda functions
- understand how file paths work
- create and delete new directories 
- know the three different types of errors
- have a plan when debugging your code