<figure>
  <IMG SRC="https://raw.githubusercontent.com/mbakker7/exploratory_computing_with_python/master/tudelft_logo.png" WIDTH=250 ALIGN="right">
</figure>

# Python Notebook #2


## Table of Contents
<ul>
    <li> <a href="#modules">2.1 Python Modules</a>
    <li> <a href="#condition_if">2.2 Conditions and if statements</a>
    <li> <a href="#datastructure">2.3 Data Structures</a>
    <li> <a href="#loops">2.4 Loops </a>
    
</ul>

<div id="modules"></div><br>
<h2>2.1 Python Modules</h2><br><div style="text-align: justify">Previously, you have learned how to: <br><br>1) initialize variables in Python; <br>2) perform simple actions with them (eg.: adding numbers together, displaying variables content, etc);<br>3) work with functions (have your own code in a function to reuse it many times or use a function, which was written by another person). <br><br>However, the scope of the last Notebook was limited by your Python knowledge and available functions in standard/vanilla Python (so-called <a href="https://docs.python.org/3/library/functions.html">'built-in' Python functions</a>).
There are not many Python built-in functions that can be useful for you, such as functions for math, plotting, signal processing, etc. Luckily, there are countless modules/packages written by other people.
</div>


<h3>Python built-in modules</h3><br><div style="text-align: justify">By installing any version of Python, you also automatically install its built-in modules.<br><br>One may wonder — why do they provide some functions within built-in modules, but not directly as a built-in function, such as the <b><code>abs()</code></b> or <b><code>print()</code></b>?<br><br>The answers may vary, but, generally, it is to keep your code clean; compact; and, working.<br><br>It keeps your code clean and compact as you only load functions that you need. It keeps your code working as it allows you to define your own functions with (almost) any kind of name, and use them easily, without worrying that you might <i>'break'</i> something if there is a function with the same name that does something completely different.</div>

<h4><code>math</code></h4><br><div style="text-align: justify">The <code>math</code> module is one of the most popular modules since it contains all implementations of basic math functions ($sin$, $cos$, $exp$, rounding, and other functions — the full list can be found <a href="https://docs.python.org/3/library/math.html">here</a>). <br><br>In order to access it, you just have to import it into your code with an <b><code>import</code></b> statement.
</div>

```python
# importing all contents of the module math

import math

print(math) # showing that math is actually a built-in Python module
```

You can now use its functions like this:

```python
print(f'Square root of 16 is equal to {int(math.sqrt(16))}')
```

You can also use the constants defined within the module, such as **`math.pi`**:

```python
print(f'π is equal to {math.pi}')

# you can also format the print() function to print the variable up to some decimal, like below:

print('π is equal to {:.2f}'.format(math.pi)) 

# change the number 2 on the ':.2f' to print with more (or fewer) decimals

print('π is equal to {:.1f}'.format(math.pi))

# this way of printing is quicker if you have to print a sentence with many variables in it, such as below:

print('π with two decimals is {:.2f}, with three decimals is {:.3f} and with four decimals is {:.4f}'.\
      format(math.pi, math.pi, math.pi))
```

<b><div class="alert alert-block alert-warning"><center>
More information on printing best practices <a href="https://realpython.com/python-string-formatting/">here</a>.</center></div></b>

<h4><code>math.pi</code></h4><br><div style="text-align: justify">As you can see, both constants and functions of a module are accessed by using: the module's name (in this case <b><code>math</code></b>) and a <b><code>.</code></b> followed by the name of the constant/function (in this case <b><code>pi</code></b>).<br><br>We are able to do this since we have loaded all contents of the module by using the <b><code>import</code></b> keyword. If we try to use these functions somehow differently — we will get an error:

```python
print('Square root of 16 is equal to')
print(sqrt(16))
```

<div style="text-align: justify">
You could, however, directly specify the functionality of the module you want to access. Then, the above cell would work.<br><br>This is done by typing: <code>from <b>module_name</b> import <b>necessary_functionality</b></code>, as shown below:

```python
from math import sqrt

print(f'Square root of 16 is equal to {int(sqrt(16))}.')
```

```python
from math import pi

print(f'π is equal to {pi}.')
```

<h4>Listing all functions</h4><br><div style="text-align: justify"> Sometimes, when you use a module for the first time, you may have no clue about the functions inside of it. In order to unveil all the potential a module has to offer, you can either access the documentation on the corresponding web resource or you can use some Python code, as shown below:
  

```python
import math

# listing all contents of a module
print('contents of math:', dir(math))

# trying to learn something about, let's say, hypot thingy
# note that \n here is used to add a new lines
print('\n\nmath hypot is a', math.hypot) 
```

```python
# you can also use ? or ?? to read the documentation about it in Python
math.hypot?
```

<h3>Python third-party modules</h3><br><div style="text-align: justify"> Besides built-in modules, there are also modules developed by other people and companies, which can be also used in your code.
These modules are not installed by default in Python, they are usually installed by using the <i>'pip'</i> or <i>'conda'</i> package managers and accessed like any other Python module.<br><br> <a href="https://www.youtube.com/watch?v=Z_Kxg-EYvxM">This YouTube video</a> explains how to install Python Packages with <i>'pip'</i> and <i>'conda'</i>. However, for the purpose of this course, you will not need to use these since Vocareum has all the needed modules built in it, without the need of accessing modules from your PC. In any case, we strongly advise you to try Python outside of Vocareum. Then, if in any doubt when installing your packages, please contact one of the TAs of this course. 

<h4><code>numpy</code></h4><br><div style="text-align: justify">The <b><code>numpy</code></b> module is one of the most popular Python modules for numerical applications. Due to its popularity, developers tend to skip using the whole module name and use a smaller version of it <b>(<code>np</code>)</b>. A different name to access a module can be done by using the <b><code>as</code></b> keyword, as shown below.

```python
import numpy as np

# printing some metainfo about the package, to show that it is not 
# a built-in Python package
print(np)

# creating 50 evenly spaced elements between 0 and 2π
x = np.linspace(0, 2 * np.pi, 50)

# calculating cos(x) for each element of x
y = np.cos(x)

# printing x and y
print('\n\nx =', x)
print('\n\ny =', y)
```

<div style="text-align: justify">You will learn more about this and other packages in separate Notebooks since these packages are frequently used by the scientific programming community.

<h4><code>matplotlib</code></h4><br><div style="text-align: justify"> The <b><code>matplotlib</code></b> module is used to plot data. It contains a lot of visualization techniques and, for simplicity, it has many submodules within the main module. Thus, in order to access functions, you have to specify the whole path that you want to import.<br><br>For example, if the function you need is located within the <b><code>matplotlib</code></b> module and <b><code>pyplot</code></b> submodule, you need to <b><code>import matplotlib.pyplot</code></b>;  then, the access command to that function is simply <b><code>pyplot.<b>your_function</b>()</code></b>.<br><br>Below we use the data generated in the previous cell to create a simple plot using the <b><code>pyplot.plot()</code></b> function:

```python
# in order to import a submodule the parent module 
# must also be specified

import matplotlib.pyplot as plt

# it is common to import matplotlib.pyplot as plt

# plt.plot() function takes x-axis values as the first argument 
# and y-axis (or f(x)) values as the second argument

plt.plot(x, y)
```

```python
# plt.scatter() works in a similar way, 
# but it does not connect the dots

plt.scatter(x, y)
```

<h3>Loading Python files as modules</h3><br><div style="text-align: justify">Finally, you can also load your own (or somebody else's) Python files as modules. This is quite helpful, as it allows you to keep your code projects well-structured without the need to copy and paste everything.<br><br>In order to import another <b>*.py</b> file as a module, you only need to have that file and your Notebook file in the same directory and use the <code>import</code> keyword. More info on this <a href='https://csatlas.com/python-import-file-module/'>here</a>.

<div class="alert alert-block alert-info"><b>(Searching) Exercise 2.1.1</b><br><br><div style="text-align: justify">Write a function, using a function from the <code>math</code> module, to find the greatest common divisor of any two numbers. Your function should return one integer number.

```python
def find_gcd(a, b):
    ...

###BEGIN SOLUTION TEMPLATE=

import math

def find_gcd(a, b):
    gcd = math.gcd(a, b)
    return gcd

###END SOLUTION
print('The greatest common divisor of 2 and 4 is:', find_gcd(2, 4))
```

<!-- grade_id: Q2.1.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert find_gcd(16, 32) == 16, '2.1.1 — Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>(Searching) Exercise 2.1.2</b><br><br><div style="text-align: justify">We can only take and store measurements at a limited number of locations and/or times. But what if we are interested in a value in between, i.e., at a location/time where we do not have a measurement? Then we can use interpolation to estimate that value. A popular, and simple, interpolation technique is <a href="https://en.wikipedia.org/wiki/Linear_interpolation">linear interpolation</a>. 
Your task is to use the module <code>scipy</code> to perform a 1D linear interpolation between a set of known points, where <code>x_known</code> and <code>y_known</code> are arrays with the measured $x$ and $y$ values. Use Google to look up the 1D interpolation function in <code>scipy</code>.<br><br>In the code below you have to replace the <code>...</code> with the correct code.

```python
... # import the module/submodule
def interpolate_inbetween(x_known, y_known, x_predict):
    f = ... # call the 1D interpolation function with proper arguments
    return f(x_predict)

###BEGIN SOLUTION TEMPLATE=
from scipy import interpolate

def interpolate_inbetween(x_known, y_known, x_predict):
    f = interpolate.interp1d(x_known, y_known)
    return f(x_predict)
###END SOLUTION

# Let's try it...
y_predict = interpolate_inbetween([1, 2], [1, 2], 1.5)
print(y_predict)
```

<!-- grade_id: Q2.1.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert abs(interpolate_inbetween([1, 2], [1, 2], 1.5) - 1.5) <= 1e-6, '2.1.2 — Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>(Searching) Exercise 2.1.3</b><br><br><div style="text-align: justify">Now, let's try to measure the running time of a function, for that we will need the <code>time</code> module. Use it to measure the working time of the <code>cool_function()</code> below.

```python
# you do not need to change anything in this cell

def cool_function():
    x = 0
    for i in range(100000000):
        x += 1

```

```python
# complete the code by replacing only the ... to create a function to calculate the running time of any function

... # remember to import the module needed

def measure_time(func):
    t0 = ... # use a function from the time module to store the current time
    func()
    t1 = ... # and again
    return ... # this should calcate the running time

###BEGIN SOLUTION TEMPLATE=
import time

def measure_time(func):
    t0 = time.time()
    func()
    t1 = time.time()
    return t1 - t0
###END SOLUTION

# Let's try it for our cool_function
print('It took {:.2f} seconds for Python to count from 0 to 100000000.'.format(measure_time(cool_function)))
```

<!-- grade_id: Q2.1.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert type(measure_time(cool_function)) == float, '2.1.3 — Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-danger"><b>Additional study material:</b>
    
* Official Python Documentation - https://docs.python.org/3/tutorial/modules.html
* https://realpython.com/python-modules-packages/
* Think Python (2nd ed.) - Sections 3 and 14 

<div id='condition_if'></div><br><h2> 2.2 Conditions and <code>if</code> statements</h2><br><div style="text-align: justify">In previous Sections you have learned how to create variables, alter them with the help of operators and access the code of professional software developers/scientists. With this, you can already do plenty of stuff in Python. However, it still lacks versatility. If you want to apply other processing techniques for other data — you would need to manually rewrite your code and then change it back once the data changes again. Not that handy, right?<br><br>In this Section you will learn how to <i>steer</i> the flow of your code — process data differently based on some conditions. For that you will learn a construction called the <b><code>if</code></b> statement.



<h3><code>if </code>keyword</h3><br><div style="text-align: justify">The <b><code>if</code></b> statement in Python is similar to how we use it in English. <i>"If I have apples, I can make an apple pie"</i> — clearly states that an apple pie will exist under the condition of you having apples. Otherwise, no pie. <br><br>
Well, it is the same in Python:

```python
amount_of_apples = 0

if amount_of_apples > 0:
    print("You have apples!\nLet's make a pie!")

print('End of the cell block...')
```

<div style="text-align: justify">As you can see - nothing is printed besides <i>'End of the cell block...'.<br><br></i>But we can clearly see that there is another print statement! Why it is not printed? Because we have no apples... thus no pie for you.<br><br>Let's acquire some fruit and see whether something will change...

```python
# adding 5 apples to our supply

amount_of_apples += 5

if amount_of_apples > 0:
    print("You have apples!\nLet's make a pie!") 

print('End of the cell block...')
```

<div style="text-align: justify">Now you can see that the same <b><code>if</code></b> statement prints text. It happened because our statement <b><code>amount_of_apples > 0</code></b> is now <b><code>True</code></b>.<br><br>That's how an <b><code>if</code></b> statement works — you type the <b><code>if</code></b> keyword, a statement and a colon. Beneath it, with an indentation of 4 spaces (1 tab), you place any code you want to run in case that <b><code>if</code></b> statement is <b><code>True</code></b>. This indentation is the same as described in Notebook 1 when defining a function.<br><br>If the result of the conditional expression is <b><code>False</code></b>, then the code inside of the <b><code>if</code></b> statement block will not run. Here's another example:

```python
my_age = 25

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(")

print('End of the cell block...')
```

<div style="text-align: justify">Slightly different setting but still the same construction. As you can see in this case, the condition of the <b><code>if</code></b> statement is more complicated than the previous one. It combines two smaller conditions by using the keyword <b><code>and</code></b>. Only if both conditions are <b><code>True</code></b> the final result is <b><code>True</code></b> (otherwise it would be <b><code>False</code></b>). <br><br>
<b><div class="alert alert-block alert-warning"><center>Thus, the condition can be as long and as complicated as you want it to be, just make sure that it is readable.</center></div></b>.

<h3><code>elif</code> keyword</h3><br>Now, let's add a bit more logic to our last example:

```python
my_age = 25

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(")
elif my_age > 65:
    print("I can finally retire!")

print('End of the cell block...')
```

Still the same output, but what if we change our age...

```python
my_age = 66

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(") # msg #1
elif my_age > 65:
    print("I can finally retire!") # msg #2

print('End of the cell block...')
```

<div style="text-align: justify">See.. we have a different output. Changing the value of our variable <b><code>my_age</code></b> changed the output of the <b><code>if</code></b> statement. Furthermore, the <b><code>elif</code></b> keyword helped us to add more logic to our code. Now, we have three different output scenarios:<br>
- print message #$1$ if <b><code>my_age</code></b> is within the $[18, 65]$ range;<br>
- print message #$2$ if <b><code>my_age</code></b> is bigger than $65$; and, <br>
- print none of them if <b><code>my_age</code></b> doesn't comply with none of the conditions (as shown below).

```python
my_age = 15

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(") # msg #1
elif my_age > 65:
    print("I can finally retire!") # msg #2

print('End of the cell block...')
```

<div style="text-align: justify">
One can also substitute an <code>elif</code> block by a different <code>if</code> block, however it is preferred to use <code>elif</code> instead to <i>"keep the condition together"</i> and to reduce code size.<br>
<div class="alert alert-block alert-warning"><center>It is important to know that there should be only <b>one</b> <code>if</code> block and <b>any number of</b> <code>elif</code> blocks within it.</center></div><br>A last example below:

```python
# Setting my_age to run the first elif block

my_age = 88

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(")
elif my_age > 65:
    print("I can finally retire!")
elif my_age < 10:
    print("I'm really, really young")

# Setting my_age to run the second elif block

my_age = 7

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(")
elif my_age > 65:
    print("I can finally retire!")
elif my_age < 10:
    print("I'm really really young")

print('End of the cell block...')
```

<h3><code>else</code> keyword</h3><br><div style="text-align: justify">We can go even further and add an additional scenario to our <b><code>if</code></b> statement with the <b><code>else</code></b> keyword. It runs the code inside of it <b>only</b> when none of the <b><code>if</code></b> and <b><code>elif</code></b> conditions are <b><code>True</code></b>:

```python
my_age = 13

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(")
elif my_age > 65:
    print("I can finally retire!")
elif my_age < 10:
    print("I'm really really young")
else:
    print("I'm just young")

print('End of the cell block...')
```

<div style="text-align: justify">
    On the previous example, since <b><code>my_age</code></b> is <b>not</b> between $[18,65]$, <b>nor</b> bigger than $65$, <b>nor</b> smaller than $10$, the <b><code>else</code></b> block is run.

Below, a final example:

```python
# Setting age to run the if block
my_age = 27

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(")
elif my_age > 65:
    print("I can finally retire!")
elif my_age < 10:
    print("I'm really really young")
else:
    print("I'm just young")

print('End of the cell block...')
print('------------------------')

# Setting age to run the first elif block
my_age = 71

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(")
elif my_age > 65: # first elif block
    print("I can finally retire!")
elif my_age < 10:
    print("I'm really really young")
else:
    print("I'm just young")

print('End of the cell block...')
print('------------------------')

# Setting age to run the second elif block
my_age = 9

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(")
elif my_age > 65:
    print("I can finally retire!")
elif my_age < 10: # second elif block
    print("I'm really really young")
else:
    print("I'm just young")

print('End of the cell block...')
print('------------------------')

# Setting age to run the else block
my_age = 13

if my_age >= 18 and my_age <= 65:
    print("I'm an adult, I have to work right now :(")
elif my_age > 65:
    print("I can finally retire!")
elif my_age < 10:
    print("I'm really really young")
else: # else block
    print("I'm just young")

print('End of the cell block...')
print('------------------------')
```

<div style="text-align: justify">That's almost everything you have to know about <b><code>if</code></b> statements! The last two things are:<br><br>
1) It goes from top to bottom. When the first condition to be <b><code>True</code></b> runs, it skips all conditions after it — as shown below:

```python
random_number = 17

if random_number > 35:
    print('Condition #1')
elif random_number > 25:
    print('Condition #2')
elif random_number > 15:
    print('Condition #3')
elif random_number > 5:
    print('Condition #4')
else:
    print('Condition #5')
```

<div style="text-align: justify">2) You can put almost everything inside each condition block and you can define variables within each block:

```python
my_income = 150
my_degree = 'BSc'

if my_degree == 'BSc':
    x = 5
    if my_income > 300:
        b = 2
        print('I am a rich BSc student')
    else:
        print('I am a poor BSc student')

elif my_degree == 'MSc':

    if my_income > 300:
        print('I am a rich MSc student')
    else:
        print('I am a poor MSc student')

print('x =', x)
print('b =', b)
```

<div style="text-align: justify">As you can see, we can make it as complicated as we want in terms of conditional branching.<br><br>Additionally, you can see that only variables within the blocks which run were created, while other variables were not. Thus, we have a <i>NameError</i> that we tried to access a variable <b><code>(b)</code></b> that was not defined.

<div class="alert alert-block alert-info"><b> Exercise 2.2.1</b><br><br><div style="text-align: justify">One of the most crucial applications of <code>if</code> statements is filtering the data from errors and checking whether an error is within a certain limit.<br><br>For example, checking whether the difference between an estimated value and the actual value are within a certain range.<br><br>
Mathematically speaking, this can be expressed as<br><br> $$|\hat{y} - y| < \epsilon$$<br>where $\hat{y}$ is your estimated value, $y$ is the actual value and $\epsilon$ is a certain error threshold.<br><br><br>The function <code>check_error_size()</code> below must do the same — it should return <code>True</code> if the error is within the acceptable range <code>eps</code> and <code>False</code> if it is not.

```python
def check_error_size(estimated_value, true_value, eps):
    ... # your if statement (error in acceptable range)
        return True
    ... # statement if not in acceptable range
        return False

###BEGIN SOLUTION TEMPLATE=
def check_error_size(estimated_value, true_value, eps):
    if abs(estimated_value - true_value) <= eps:
        return True
    else:
        return False
###END SOLUTION

#You can try to check it by yourself by running the function with some self-chosen numbers
check_error(...)
```

<!-- grade_id: Q2.2.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert check_error(0.5, 0.4, 0.2), '2.2.1 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>Exercise 2.2.2</b><br><br><div style="text-align: justify">Use the knowledge you have obtained in this Notebook and write your very own function to classify soil samples based on the average grain size. The classification you have to implement is the following:<br>

1.   Clay: avg grain size $<$ 0.002 mm
2.   Silt: 0.002 mm $\leq$ avg grain size $<$ 0.063 mm
3.   Sand: 0.063 mm $\leq$ avg grain size $<$ 2 mm
4.   Gravel: 2 mm $\leq$ avg grain size $<$ 63 mm

<div style="text-align: justify">Your task is to write the function, which will return the name of the soil type based on the provided average grain size in meters.

```python
def classify_soil(avg_grain_size):
    ...

###BEGIN SOLUTION TEMPLATE=
def classify_soil(avg_grain_size):
    if avg_grain_size < 0.002:
        return 'Clay'
    elif avg_grain_size >= 0.002 and avg_grain_size < 0.063:
        return 'Silt'
    elif avg_grain_size >= 0.063 and avg_grain_size < 2:
        return 'Sand'
    elif avg_grain_size >= 2:
        return 'Gravel'
###END SOLUTION

print(classify_soil(1.5))
```

<!-- grade_id: Q2.2.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert classify_soil(0.0015) == 'Clay' and \
classify_soil(0.05) == 'Silt' and \
classify_soil(0.1) == 'Sand' and \
classify_soil(2.5) == 'Gravel' , '2.2.2 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>Exercise 2.2.3 — Triangle Inequality</b><br><br><div style="text-align: justify">Let's imagine that you received a request to manage a web-application, which provides and visualizes <a href="https://en.wikipedia.org/wiki/Interferometric_synthetic-aperture_radar"><b>InSAR</b></a> data of the estimated displacement in an area.<br><br>The app is already quite cool, however, there is always room for improvement. For example, one might want to look into the statistics over the area of a triangle.<br><br>Your goal is to help implement such functionality by checking whether the user selected a valid triangle. You need to do this by checking that the length of <b>each</b> of the sides is smaller than the sum of the other two.

```python
def check_triangle(side_a, side_b, side_c):
    ...

###BEGIN SOLUTION TEMPLATE=
def check_triangle(side_a, side_b, side_c):
    part1 = side_a < (side_b + side_c)
    part2 = side_b < (side_a + side_c)
    part3 = side_c < (side_a + side_b)
    

    if part1 and part2 and part3:
        return "Valid Triangle"
    else:
        return "Invalid Triangle"
###END SOLUTION
```

<!-- grade_id: Q2.2.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert check_triangle(1, 2, 3) == "Invalid Triangle" and check_triangle(3, 4, 5) == "Valid Triangle", '2.2.3 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>(Searching) Exercise 2.2.4</b><br><br><div style="text-align: justify">Conditional expression is a way how one can compress an <code>if</code> statement to a more compact (and logical) statement. Your task is to rewrite the below <code>if</code> statement by using the <i>'conditional expression'</i> technique.

```python
y = 0
x = 5

if x % 2 == 0:
    y = x ** 2
else:
    y = x % 3
    

```

```python
# write your code here
y = ...

###BEGIN SOLUTION TEMPLATE=
y = x ** 2 if x % 2 == 0 else x % 3
###END SOLUTION

print(y)
```

<div class="alert alert-block alert-danger"><b>Additional study material:</b>
    
* Official Python Documentation - https://docs.python.org/3/tutorial/controlflow.html
* https://realpython.com/python-conditional-statements/
* Think Python (2nd ed.) - Section 5

<div id="datastructure"></div><br><h2>2.3 Data Structures</h2><br><div style="text-align: justify">In this Section you will tackle a data management problem! In the first module you have learned how to create variables, which is cool. But when you populate a lot of variables, or you want to store & access them within one entity, you need to have a data structure.<br><br>There are plenty of them, which differ their use cases and complexity. Today we will tackle some of the standard Python built-in data structures. The most popular of those are: <b><code>list</code></b>, <b><code>dict</code></b> and <b><code>tuple</code></b>.

<h3><code>list</code></h3><br><div style="text-align: justify">First, the easiest and the most popular data structure in Python: <b><code>list</code></b> (which is similar to a typical array you could have seen in a different programming language).<br><br>
You can create a list in the following ways:

```python
# 1). Creating an empty list, option 1

empty_list1 = []
print('Type of my_list1 object', type(empty_list1))
print('Contents of my_list1', empty_list1)
print('--------------------')

# 2). Creating an empty list, option 2 - using the class constructor

empty_list2 = list()
print('Type of my_list2 object', type(empty_list2))
print('Contents of my_list2', empty_list2)
print('--------------------')

# 3). Creating a list from existing data - option 1

my_var1 = 5
my_var2 = "hello"
my_var3 = 37.5

my_list = [my_var1, my_var2, my_var3]
print('Type of my_list3 object', type(my_list))
print('Contents of my_list3', my_list)
print('--------------------')

# 4). Creating a list from existing data - option 2

cool_rock = "sandstone" # remember that a string is a collection of characters

list_with_letters = list(cool_rock)

print('Type of my_list3 object', type(list_with_letters))
print('Contents of list_with_letters', list_with_letters)
print('--------------------')

```

<div style="text-align: justify">As you can see, in all three cases we created a list, only the method how we did it was slightly different: <br><br>
- the first method uses the bracket notation,<br>
- the second method uses class constructor approach. <br><br>
Both methods also apply to the other data structures.<br><br>
Now, we have a list — what can we do with it?
Well... we can access and modify any element of an existing list. In order to access a list element, square brackets <b><code>[]</code></b> are used with the index of the element we want to access inside. Sounds easy, but keep in mind that Python has a zero-based indexing (as mentioned in Section 1.4 in Notebook 1).<br><br> 
  As a reminder, a zero-based indexing means that the first element has index 0 (not 1), the second element has index 1 (not 2) and the n-th element has index n - 1 (not n)!

```python
# len() function returns the lengths of an iterable (string, list, array, etc)
print(len(my_list))

# We have 3 elements, thus we can access 0th, 1st, and 2nd elements
print('First element of my list:', my_list[0])

print('Last element of my list:', my_list[2])

# After the element is accessed, it can be used as any variable,
# the list only provides a convenient storage

summation = my_list[0] + my_list[2]
print(f'Sum of {my_list[0]} and {my_list[2]} is {summation}')

# Since it is a storage - we can easily alter and swap list elements
my_list[0] += 7
my_list[1] = "My new element"

print(my_list)

# However we can only access data we have - Python will give us an error for the following

my_list[10] = 199
```

<div style="text-align: justify">
We can also add new elements to a list, or remove them! Adding is realized with the <b><code>append</code></b> method and removal of an element uses the <b><code>del</code></b> keyword.

```python
# adding a new element to the end of the list
my_list.append("new addition to  my variable collection!")
print(my_list)

# we can also store a list inside a list - list inception! Useful for matrices, images etc
my_list.append(['another list', False, 1 + 2j])
print(my_list)

# Let's remove 37.5

del my_list[2]

print(my_list)
```

<div style="text-align: justify">
Lists also have other useful functionalities, as you can see from the <a href="https://docs.python.org/3/tutorial/datastructures.html">official documentation</a>. Since lists are still objects you can try and apply some operations to them as well.

```python
lst1 = [2, 4, False]
lst2 = ['second list', 0, 222]

#what will happen?
lst1 = lst1 + lst2
print(lst1)

lst2 = lst2 * 4
print(lst2)

lst2[3] = 5050
print(lst2)
```

<div style="text-align: justify">
As you can see, adding lists together concatenates them and multiplying them basically does the same thing (it performs addition several times, just like in real math...).<br><br> Additionally, you can also use the <b><code>in</code></b> keyword to check the presence of a value inside a list.

```python
print(lst1)

if 222 in lst1:
    print('We found 222 inside lst1')
else:
    print('Nope, nothing there....')
```

<h3><code>tuple</code></h3><br><div style="text-align: justify">If you understood how <code>list</code> works, then you already understand 95% of <b><code>tuple</code></b>. Tuples are just like lists, with some small differences.<br><br>1.   In order to create a tuple you need to use <b><code>()</code></b> brackets, comma or a <b><code>tuple</code></b> class constructor.<br>2.   You can change the content of your list, however tuples are immutable (just like strings).



```python
# Creating an empty tuple - a bit useless, since you cannot change it

tupl1 = tuple() # option 1 with class constructor
print('Type of tupl1', type(tupl1))
print('Content of tupl1', tupl1)

tupl2 = () # option 2 with ()
print(type(tupl2), tupl2)
```

```python
# Creating a non-empty tuple using brackets

my_var1 = 26.5
my_var2 = 'Oil'
my_var3 = False

my_tuple = (my_var1, my_var2, my_var3, 'some additional stuff', 777)
print('my tuple', my_tuple)

# Creating a non-empty tuple using comma

comma_tuple = 2, 'hi!', 228
print('A comma made tuple', comma_tuple)

# now, let's try to access an element
print('4th element of my_tuple:', my_tuple[3])

# but, can we change it?
my_tuple[3] = 'will I change?'
```

<div style="text-align: justify">
Since tuples are immutable, it has no <b><code>append()</code></b> method nor any other methods that alter the object they target.<br><br>
You might think that tuple is a useless class. However, there are some reasons for it to exist:<br><br>
1)   Storing constants & objects which shouldn't be changed.<br>
2)   Saving memory (tuple uses less memory to store the same data than a list).



```python
#creating a list and a tuple from the same data

my_name = 'Vasyan'
my_age = 27
is_student = True

a = (my_name, my_age, is_student)
b = [my_name, my_age, is_student]

print('size of a =', a.__sizeof__(), 'bytes') #.__sizeof__() determines the size of a variable in bytes
print('size of b =', b.__sizeof__(), 'bytes')
```

<h3><code>dict</code></h3><br><div style="text-align: justify">After seeing lists and tuples, you may think: <br><br>"Wow, storing all my variables within another variable is cool and gnarly! But... sometimes it's boring & inconvenient to access my data by using it's position within a tuple/list. Is there a way that I can store my object within a data structure but access it via something meaningful, like a keyword...?"<br><br>Don't worry if you had this exact same thought.. Python had it as well!<br><br>Dictionaries are suited especially for that purpose — to each element you want to store, you give it a nickname (i.e., a key) and use that key to access the value you want.

```python
# Creating an empty dictionary - we used () for tuples and and [] for lists. 
# Now, it's time to use {}.

empty_dict1 = {}
print('Type of empty_dict1', type(empty_dict1))
print('Content of it ->', empty_dict1)

# Creating an empty dictionary - using class constructor
empty_dict2 = dict()
print('Type of empty_dict2', type(empty_dict2))
print('Content of it ->', empty_dict2)
```

```python
# Creating a non-empty dictionary - specifying pairs of key:value pattern
my_dict = {
    'name': 'Jarno',
    'color': 'red',
    'year': 2007,
    'is cool': True,
    6: 'it works',
    (2, 22): 'that is a strange key'
}

print('Content of my_dict>>>', my_dict)
```

<div style="text-align: justify">
In the last example, you can see that only strings, numbers, or tuples were used as keys. Dictionaries can only use immutable data (or numbers) as keys:

```python
# using mutable structures won't work: 

mutable_key_dict = {
    5: 'lets try',
    True: 'I hope it will run perfectly',
    6.78: 'heh',
    ['No problemo', 'right?']: False  
}

print(mutable_key_dict)
```

<div style="text-align: justify">
Alright, now it is time to access the data we have managed to store inside <b><code>my_dict</code></b>...

```python
# for that we use keys!

print('Some random content of my_dict', my_dict['name'], my_dict[(2, 22)])
```

```python
# remember the mutable key dict? Let's make it work by omitting the list item
mutable_key_dict = {
    5: 'lets try',
    True: 'I hope it will run perfectly',
    6.78: 'heh'
}

# You can see that it doesn't give any errors but how do we access the data inside it?
# use keys!
print('Accessing weird dictionary...')
print(mutable_key_dict[True])
print(mutable_key_dict[5])
print(mutable_key_dict[6.78])
```

```python
# Trying to access something we have and something we don't have
print('My favorite year is', my_dict['year'])
print('My favorite song is', my_dict['song'])
```

<div class="alert alert-block alert-warning"><div style="text-align: justify"><center>It is best practice to use mainly <i><b>strings</b></i> as keys — the other options are weird and are almost never used.</center></div></div><br><div style="text-align: justify">What's next? Dictionaries are mutable, so let's go ahead and add some additional data and delete old ones.

```python
print('my_dict right now', my_dict)

my_dict['new_element'] = 'magenta'
my_dict['weight'] = 27.8
del my_dict['year']

print('my_dict after some operations', my_dict)
```

<div style="text-align: justify">
    You can also print all keys present in the dictionary using the <b><code>.keys()</code></b> method, or check whether a certain key exists in a dictionary, as shown below. More operations can be found <a href="https://docs.python.org/3/tutorial/datastructures.html">here</a>.

```python
print(my_dict.keys())

# check if my_dict has a name key
print("\nmy_dict has a ['name'] key:", 'name' in my_dict)
```

<div class="alert alert-block alert-success"><center><b>Real life example: Analyzing satellite metadata<br><br></b>Metadata is a set of data that describes and gives information about other data. For Sentinel-1, the metadata of the satellite is acquired as an <i>.xml</i> file. It is common for Dictionaries to play an important role in classifying this metadata. One could write a function to read and obtain important information from this metadata and store them in a Dictionary. Some examples of keys for the metadata of Sentinel-1 are:</center><br>
    
<i><center>dict_keys(['azimuthSteeringRate', 'dataDcPolynomial', 'dcAzimuthtime', 'dcT0', 'rangePixelSpacing', 'azimuthPixelSpacing', 'azimuthFmRatePolynomial', 'azimuthFmRateTime', 'azimuthFmRateT0', 'radarFrequency', 'velocity', 'velocityTime', 'linesPerBurst', 'azimuthTimeInterval', 'rangeSamplingRate', 'slantRangeTime', 'samplesPerBurst', 'no_burst'])</i>


<h3>Slices</h3><br><div style="text-align: justify">The last important thing for this Notebook are slices. Similar to how you can slice a <b>string</b> (shown in Section 1.4, in Notebook 1). This technique allows you to select a subset of data from an iterable (like a list or a tuple).

```python
# let's make a simple list

x = [1, 2, 3, 4, 5, 6, 7]
n = len(x) # len(x) gives the length of x

# in order to select a slice of it, you have to use a : symbol,
# like this => my_list[start:end], 
# it will select all elements with indices [start, end), 
# starting from start and ending at end-1 (excluding the end).
print('The first three elements of x:', x[0:3])

# getting the same subset but using a different slice "equation"

# if we start from the beginning (at index 0) - we can omit the 0
print(x[:3])

# instead of counting elements from the beginning, we can count from end
print('The last element is', x[6], 'or', x[n - 1], 'or', x[-1])

# thus, we can apply it in slicing our list
print(x[0:-4])

# we can also specify a third argument: the step size of our selection
print(x[0:3:1])
```

<div style="text-align: justify">
Thus, the general slicing call is given by <b><code>iterable[start:end:step]</code></b>. <br><br>Here's another example:

```python
numbers = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]

# similar to omiting the 0 if you start at the beginning of the list
# you can omit the last value if you want all values until the end
print('Selecting all even numbers', numbers[::2])
# in the above case you start at 0 (omited), end at 10 (omited)
# with a step of 2 (not omited)

print('All odd numbers', numbers[1::2])

# Reversing the list
print('Normal order', numbers)
print('Reversed order', numbers[::-1])

# Selecting middle subset
print('Numbers from 5 to 8:', numbers[5:9])
```

<div class="alert alert-block alert-info"><b>Exercise 2.3.1</b><br><br><div style="text-align: justify">Now, let's get down to practice.<br><br>Your first task is to finish the <code>pack_variables()</code> function — a function which will combine all inputs in one list and return it.<br><br>More precisely, this function will receive $5$ arguments as input and you have to return a list with these $5$ elements inside.

```python
def pack_variables(arg1, arg2, arg3, arg4, arg5):
    ...

###BEGIN SOLUTION TEMPLATE=
def pack_variables(arg1, arg2, arg3, arg4, arg5):
    pack = [arg1, arg2, arg3, arg4, arg5]
    return pack
###END SOLUTION

print(pack_variables(1, 2, 4, 22, 7))

```

<!-- grade_id: Q2.3.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert len(pack_variables(221, 22, "2", 22, 7)) == 5 and type(pack_variables(221, 22, "2", 22, 7)) == str, '2.3.1 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>Exercise 2.3.2</b><br><br><div style="text-align: justify">Here, you will have to perform a quality assessment on a received list of GNSS measurements. But a simple one.<br><br>You have to check whether the received data has more than $1000$ measurements, in that case we know the receiver was running for a long time without being interrupted.
If it is shorter than $1000$, we assume there were some interruptions. Hence, your function should return a message whether the data is fine or not.

```python
def check_data(measurements):
    ...

###BEGIN SOLUTION TEMPLATE=
def check_data(measurements):
    if len(measurements) < 1000:
        return "Something wrong"
    else:
        return "All fine"
###END SOLUTION

#You can check your code below, where we simulate 1250 measurements
import random 

gnss_data = [random.random() * 2.2e8 for i in range(1250)]
print(check_data(gnss_data))
```

<!-- grade_id: Q2.3.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert check_data([1,2,3]) == "Something wrong" and check_data([0 for i in range(1001)]), '2.1.2 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>Exercise 2.3.3</b><br><br><div style="text-align: justify">In this exercises you need to write something opposite to a packager... an unpacker! Sometimes it is easier to work with the most convenient data type, that is best suited for a specific task. In this example, you have to write a function which accepts data stored in a dictionary and saves some data from it. More precisely, you have to select $x$ and $y$ coordinates saved in the <code>input_dict</code> dictionary, and return them as a tuple, with $x$ being the first entry.

```python
# you do not have to change anything in this cell

input_dict = {
    'ID': '334856',
    'operator_name': 'Jarno',
    'observation_amount': '2485',
    'loc_x': [2, 5, 10, 12, -5, 8, 27, 1],
    'loc_y': [6, 1, -5, 15, 4, 8, 0, 10]
}

```

```python
def unpack_dictionary(input_dict):
    ...

###BEGIN SOLUTION TEMPLATE=
def unpack_dictionary(input_dict):
    return (input_dict['loc_x'], input_dict['loc_y'])
###END SOLUTION    

print(unpack_dictionary(input_dict))
```

<!-- grade_id: Q2.3.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert type(unpack_dictionary(input_dict)) == tuple, '2.3.3 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>(Searching) Exercise 2.3.4 — List comprehension</b><br><br><div style="text-align: justify">Sometimes you have to create a list of predefined size and fill it with data. This can be done easily using list comprehension.
Write a function using list comprehension, in order to create a list of a predefined size and fill them with zeros.

```python
def create_list(size):
    ...

###BEGIN SOLUTION TEMPLATE=
def create_list(size):
    new_list = [0 for i in range(size)]
    return new_list
###END SOLUTION

print(create_list(5))
```

<!-- grade_id: Q2.3.4, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert len(create_list(5)) == 5 and type(create_list(5)) == list, '2.3.4 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-danger"><b>Additional study material:</b>
    
* Official Python Documentation - https://docs.python.org/3/tutorial/datastructures.html
* https://realpython.com/python-data-structures/
* Think Python (2nd ed.) - Sections 8, 10, 11, 12

<div id='loops'></div><br><h2>2.4 Loops</h2><br><div style="text-align: justify">Let's do another step to automatize things even more! Previous Sections introduced a lot of fundamental concepts, but they still don't unveil the true power of any programming language — loops!<br><br>If we want to perform the same procedure multiple times, then we would have to take the same code and copy-paste it. This approach would work, however it would require a lot of manual work and it does not look cool.<br><br>This problem is resolved with a <i>loop</i> construction. As the name suggest, this construction allows you to loop (or run) certain piece of code several times at one execution.

<h3><code>for</code> loop</h3><br><div style="text-align: justify">The first and the most popular looping technique is a <b><code>for</code></b> loop. Let's see some examples:

```python
# let's create a list with some stuff in it

my_list = [100, 'marble', False, 2, 2, [7, 7, 7], 'end']

# in order to iterate (or go through each element of a list)
# we use a for loop

print('Start of the loop')
for list_item in my_list:
    print('In my list I can find:', list_item)
print('End of the loop')
```

<div style="text-align: justify">
General <code>for</code> loop construction looks like this:<br><br>
<code>for iterator_variable in iterable:
       do something with iterator_variable</code><br><br>
During each iteration the following steps are happening under the hood (or above it):<br><br>1.   <b><code>iterator_variable = iterable[0]</code></b> iterator_variable is assigned the first value from the iterable.<br>2.   Then, you use iterator_variable as you wish<br>3.   By the end of the 'cycle', the next element from the iterable is selected <b>(<code>iterable[1]</code>)</b>, i.e., we return to step 1, but now assigning the second element... and so on.<br>4. When there is not a next element (in other words, we have reached the end of the iterable) — it exits and the code under the loop is now executed.<br><br>Looks cool, but what if we want to alter the original iterable (not the <b><code>iterator_variable</code>)</b> within the loop?




```python
# let's see whether we can change the original list within a list 
# with a for loop

x = my_list
print('Try #1, before:', x)

for item in x:
    item = [5,6,7]

print('Try #1, after', x)

```

Nothing has changed.... let's try another method.

```python
length_of_x = len(x)

# range() is used to generate a sequence of numbers
# more info at https://www.w3schools.com/python/ref_func_range.asp

indices = range(length_of_x) # this will generate numbers from 0 till length_of_x, excluding the last one

print(indices)
print('Try #2, before', my_list)

for id in indices:
    my_list[id] = -1

print('Try #2, after', my_list)
 

```

<div style="text-align: justify">Now we have a method in our arsenal which can not only loop through a list but also access and alter its contents. Also, you can generate new data by using a <b><code>for</code></b> loop and by applying some processing to it. Here's an example on how you can automatize your greetings routine!

```python
# General greeting
msg = "Ohayo"

# the list with your friends names
names = ["Jarno", "Alex", "John", "Maria", "Xenia", "Janis", "Vasya"]

# An empty list, where all greetings will be stored (otherwise you cannot use the .append in the for loop below!)
greetings = []

for name in names:
    personalized_greeting = f'{msg}, {name}-kun!' # create the personalize greeting
    greetings.append(personalized_greeting) # append the personalized meeting tot the list of greetings

# Printing our newly created greetings
print(greetings)
```

And you can also have loops inside loops!

```python
# Let's say that you put down all your expenses per day separately, 
# in euros
day1_expenses = [15, 100, 9]
day2_expenses = [200]
day3_expenses = [10, 12, 15, 5, 1]

# you can also keep them within one list together
expenses = [day1_expenses, day2_expenses, day3_expenses]
print('All my expenses', expenses)

# you can access also each expense separately!
# day3 is third array and 2nd expense is second element 
# within that array
print(f'My second expense on day 3 is {expenses[2][1]}')# recall 0th based indexing!
```

```python
# Now let's use it in some calculations 

# Option #1
total_expenses = 0

for i in range(len(expenses)): # loop over all days
    # Accessing expenses made at day i + 1
    daily_expenses_list = expenses[i]
    # Creating temporary storage for current day expenses
    daily_expenses = 0
    
    for j in range(len(daily_expenses_list)): # loop over expenses for day i + 1
        daily_expenses += daily_expenses_list[j]
    
    # Adding daily expenses to the total expenses
    total_expenses += daily_expenses
    
print(f'Option #1: In total I have spent {total_expenses} euro!')
```

```python
# Option #2
total_expenses = 0

for i in range(len(expenses)):
    for j in range(len(expenses[i])):
        total_expenses += expenses[i][j]
    
print(f'Option #2: In total I have spent {total_expenses} euro!')
```

```python
# Option #3 - advanced techniques gathered after eternal suffering
total_expenses = 0
total_expenses = sum(map(sum, expenses))
print(f'Option #3: In total I have spent {total_expenses} euro!')
```

<h3><code>while</code> loop</h3><br><div style="text-align: justify">The second popular loop construction is a <b><code>while</code></b> loop. The main difference is that it is suited for code structures that must repeat unless a certain logical condition is satisfied. It looks like this:<br><br><center><code>while logical_condition == True:<br>do something</code></center><br>And here is a working code example:


```python
sum = 0

while sum < 5:
    print('sum in the beginning of the cycle:', sum)
    sum += 1
    print('sum in the end of the cycle:', sum)
```

<div style="text-align: justify">As you can see, this loop was used to increase the value of the sum variable until it reached $5$. The moment it reached $5$ and the loop condition was checked — it returned <b><code>False</code></b> and, therefore, the loop stopped.<br><br>Additionally, it is worth to mention that the code inside the loop was altering the variable used in the loop condition statement, which allowed it to first run, and then stop. In the case where the code doesn't alter the loop condition, it won't stop (infinite loop), unless another special word is used.<br><br>Here's a simple example of an infinite loop, which you may run (by removing the #'s) but in order to stop it — you have to interrupt the Notebook's kernel or restart it.

```python
# a, b = 0, 7

# while a + b < 10:
#     a += 1
#     b -= 1
#     print(f'a:{a};b:{b}')
```

<div class="alert alert-block alert-warning"><div style="text-align: justify"><center>Before submitting your notebook in Vocareum, make sure to comment out the above code (or any other code with infinite loops) with <b>#</b>'s again. You can select the entire code block and press <code>Ctrl+/</code> to comment out everything.

<h3><code>break</code> keyword</h3><br><div style="text-align: justify">After meeting and understanding the loop constructions, we can add a bit more control to it. For example, it would be nice to exit a loop earlier than it ends — in order to avoid infinite loops or just in case there is no need to run the loop further. This can be achieved by using the <b><code>break</code></b> keyword. The moment this keyword is executed, the code exits from the current loop.

```python
stop_iteration = 4

print('Before normal loop')
for i in range(7):
    print(f'{i} iteration and still running...')
print('After normal loop')

print('Before interrupted loop')
for i in range(7):
    print(f'{i} iteration and still running...')

    if i == stop_iteration:
        print('Leaving the loop')
        break
print('After interupted loop')
```

<div style="text-align: justify">
    The second loop shows how a small intrusion of an <b><code>if</code></b> statement and the <b><code>break</code></b> keyword can help us with stopping the loop earlier. The same word can be also used in a <b><code>while</code></b> loop:

```python
iteration_number = 0

print('Before the loop')
while True:
    iteration_number += 1

    print(f'Inside the loop #{iteration_number}')
    if iteration_number > 5:
        print('Too many iterations is bad for your health')
        break
print('After the loop')
```

<h3><code>continue</code> keyword</h3><br><div style="text-align: justify">Another possibility to be more flexible when using loops is to use the <b><code>continue</code></b> keyword.<br><br>This will allow you to skip some iterations (more precisely — the moment the keyword is used it will skip the code underneath it and will start the next iteration from the beginning).

```python
def calculate_cool_function(arg):
    res = 7 * arg ** 2 + 5 * arg + 3
    print(f'Calculating cool function for {arg} ->  f({arg}) = {res}')

print('Begin normal loop\n')
for i in range(7):
    print(f'{i} iteration and still running...')
    calculate_cool_function(i)
print('\nEnd normal loop\n')

print('-------------------')

print('Begin altered loop\n')
for i in range(7):
    print(f'{i} iteration and still running...')

    # skipping every even iteration
    if i % 2 == 0:
        continue
        
    calculate_cool_function(i)
    
print('\nEnd altered loop')
```

<div style="text-align: justify">As you can see, with the help of the <b><code>continue</code></b> keyword we managed to skip some of the iterations. Also worth noting that $0$ is divisible by any number, for that reason the <b><code>calculate_cool_function(i)</code></b> at <b><code>i = 0</code></b> didn't run.

<div class="alert alert-block alert-info"><b>Exercise 2.4.1</b><br><br><div style="text-align: justify">In this exercise you will write your own Celsius to Fahrenheit converter! Your task is to write a function, which will accept the list of temperatures <code>temp_c</code>, in Celsius, and will output a list with the same temperatures, but in Fahrenheit.

```python
# you do not need to change anything in this cell

temp_c = [-1, -1.2, 1.3, 6.4, 11.2, 14.8, 17.8, 17.7, 13.7, 8.5, 4.1, 0.9]
```

```python
def celsius_to_fahrenheit(temp_c):
    ...

###BEGIN SOLUTION TEMPLATE=
def celsius_to_fahrenheit(temp_c):
    temp_f = []

    for i in range(len(temp_c)):
        temp_f.append(temp_c[i] * 9 / 5 + 32) 
    return temp_f
###END SOLUTION

print(celsius_to_fahrenheit(temp_c))
```

<!-- grade_id: Q2.4.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert type(celsius_to_fahrenheit([-1])) == list and abs(celsius_to_fahrenheit([-1])[0] - 30.2) <= 1e-6, '2.4.1 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>Exercise 2.4.2</b><br><br><div style="text-align: justify">Your task here is to write a function which will analyze a broadcasting message of the following format: <b>"satellite_ids;date"</b>, where the first part of the message contains unique lowercase letters, each corresponding to a different satellite ID, and the last part contains the date of the message.<br><br>
Here are some examples:
    <b>"agf;06062022"</b>,<b> "abcdefgops;03121999" </b>or<b> "xyz;11112011".</b>
<br><br>Your task is to write a function, which for a provided broadcast message, will count the number of satellites mentioned in the message.


```python
def count_satellites(message):
    ...

###BEGIN SOLUTION TEMPLATE=
def count_satellites(message):
    for i in range(len(message)):
        if message[i] == ';':
            break

    return i
###END SOLUTION

# Check that with the example below you count 5 satellites
print(count_satellites("hallo;12122007"))

```

<!-- grade_id: Q2.4.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert count_satellites("qute;12122007") == 4, '2.4.2 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>Exercise 2.4.3</b><br><br><div style="text-align: justify">Here you need to write a function that is able to sort any list consisting only of real numbers, in the descending order. For example, the list $[19, 5, 144, 6]$ becomes $[144, 19, 6, 5]$. Hint: use a built-in function or write your own bubble-sort algorithm to sort the list in ascending order, and then think of a clever way to change the order this list to descending order.

```python
def sort_list(unsorted_list):
    ...

###BEGIN SOLUTION TEMPLATE=
def sort_list(unsorted_list):
    return sorted(unsorted_list)[::-1]
###END SOLUTION


print(sort_list([9, -1, 5, 1, -9, -9]))
```

<!-- grade_id: Q2.4.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert sort_list([3, 1, 2]) == [3, 2, 1], '2.4.3 - Incorrect answer'
###END HIDDEN TESTS
```

<div style="text-align: justify">Before you continue with exercises 2.4.4 and 2.4.5 we would like to try and facilitate your visualization of how to 'navigate' a 2D object such as a Matrix or a 2D list. All comes down to understanding indexing.

```python
import numpy as np # ignore the use of np.array at this point, it just facilitates the printed
# 'matrix view' below

mat = np.array([['[0][0]','[0][1]','[0][2]','[0][3]'],['[1][0]','[1][1]','[1][2]','[1][3]'],['[2][0]','[2][1]','[2][2]','[2][3]']])
print(mat)
```

Output:
```
[['[0][0]' '[0][1]' '[0][2]' '[0][3]']
 ['[1][0]' '[1][1]' '[1][2]' '[1][3]']
 ['[2][0]' '[2][1]' '[2][2]' '[2][3]']]
```

<div style="text-align: justify">To facilitate this illustration, the above Matrix has its own index as elements. In a 2D object, the first <b>indexing</b> is the number of the <b>row</b>, the second is the number of the <b>column</b>. Let's break it down.

```python
print(mat[0]) # the 0-th row of the variable 'mat'
```

Output:
```
['[0][0]' '[0][1]' '[0][2]' '[0][3]']
```

```python
print(mat[1]) # the 1-st row of the variable 'mat'; etc..
```

Output:
```
['[1][0]' '[1][1]' '[1][2]' '[1][3]']
```

```python
print(mat[2,3]) # the 2-nd row and 3-th column
```

Output:
```
[2][3]
```

```python
print(mat[:,2]) # all rows and 2-nd column (in other words, only the 2-nd column)
```

Output:
```
['[0][2]' '[1][2]' '[2][2]']
```

```python
print(mat[-1,:]) # last row, all columns (therefore, last row)
```

Output:
```
['[2][0]' '[2][1]' '[2][2]' '[2][3]']
```

```python
# could have also been achieved only by writing mat[-1]
print(mat[-1])
```

Output:
```
['[2][0]' '[2][1]' '[2][2]' '[2][3]']
```

```python
# check the number of rows of the var 'mat'
print(len(mat))
```

Output:
```
3
```

```python
# check the number of columns of the var 'mat'
print(len(mat[0])) # number of elements in row [0]
```

Output:
```
4
```

Note that here we only check the number of elements in the 0-th row, since all rows have same number of elements. However, you may also work with irregular lists, where each 'row' or 'sublist' will have a different number of 'columns'.

Now that you are a bit more familiarized with indexing, let's hop into the last two exercises of this Notebook.

<div class="alert alert-block alert-info"><b>Exercise 2.4.4</b><br><br><div style="text-align: justify">Quite frequently you will have to work with 2D data, which (in Python) is stored, sometimes, in 2 dimensional lists. Your task is to complete a function, which will count elements in an irregular 2D list (irregular means that the amount of elements within each sub-list is different). For that you will have to write a loop inside a loop (or, in other words, a double loop). The code below is almost complete, you just need finish typing the second <code>for</code> loop.

```python
def count_elements(data):
    #initializing counter
    cnt = 0    
    for i in range(len(data)) # looping through the number of rows
        for ... # looping to count elements in each row (or sub-list, in this case)
            cnt += 1
            
    return cnt

###BEGIN SOLUTION TEMPLATE=
def count_elements(data):
    #initializing counter
    cnt = 0    
    for i in range(len(data)):
        for j in range(len(data[i])):
            cnt += 1
            
    return cnt

###END SOLUTION

test = [[1, 2, 3], ["list", 4.5], [9, -1, "another list", 6, 101]]
print(f'Total number of elements inside a test list: {count_elements(test)}')
```

Output:
```
Total amount of elements inside a test list: 10
```

<!-- grade_id: Q2.4.4, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
test_data = [[1], [2, 3], [4, 5, 6], [7, 8, "9", 10]]
assert count_elements(test_data) == 10, '2.4.4 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-info"><b>Exercise 2.4.5</b><br><br><div style="text-align: justify">In this exercise you will perform <a href="https://en.wikipedia.org/wiki/Downsampling_(signal_processing)">downsampling</a> of a provided 'regular' 2D list. Downsampling is a procedure where only a subset of the data is sampled (to reduce its size, for example). Below a visual aid of what downsampling of a 2D list is. Instead of using all the data available, your task is to downsample the 2D list to keep only the data of every $a$-th row and every $b$-th column, including the first element $(a_{0,0})$.<br><br>$$A = \left[\begin{array}{ccccc}
a_{0,0} & \dots & a_{0,b} & \dots & a_{0,2b} & \dots & \dots & a_{0,nb}	& \dots\\
\vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots\\
a_{a,0} & \dots & a_{a,b} & \dots & a_{a,2b} & \dots & \dots & a_{a,nb}	& \dots\\
\vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots\\
a_{2a,0} & \dots & a_{2a,b} & \dots & a_{2a,2b} & \dots & \dots & a_{2a,nb}	& \dots\\
\vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots\\
\vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots\\
a_{ma,0} & \dots & a_{ma,b} & \dots & a_{ma,2b} & \dots & \dots & a_{ma,nb} & \dots \\
\vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots
\end{array}\right]$$
<br><i>Hint: Use the index value of each element to realize if they should be included or not in the downsample_data list.<br>Hint2: Remember that $0$ % $X = 0$, for any $X$. (Modulus operator, discussed in Notebook 1).

```python
def downsample(data, a, b):
    downsample_data = []

    # create a for loop over len(data), which is the number of rows
    for ...
        row = []

        # create a for loop over len(data[i]), which is the number of columns
        for ...
            # if an element is to be stored, append row variable (created above)
            ... 
            ...

        # finally append downsample_data      
        if len(row) > 0:
            ...
         
    return ...

print(downsample(insar_data, 3, 4))
###BEGIN SOLUTION TEMPLATE=
def downsample(data, a, b):
    downsample_data = []
    
    for i in range(len(data)):
        row = []
        for j in range(len(data[i])):
            if i % a == 0 and j % b == 0:
                row.append(data[i][j])
        if len(row) > 0:
            downsample_data.append(row)
            
    return downsample_data
###END SOLUTION

import numpy as np

insar_data = [[np.random.rand() for i in range(12)] for i in range(9)]

print(downsample(insar_data, 3, 4))
```

<!-- grade_id: Q2.4.5, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
insar_data = [[np.random.rand() for i in range(16)] for i in range(12)]
downsampled = downsample(insar_data, 3, 4)
dim1 = len(downsampled)
dim2 = len(downsampled[0])
assert dim1 == 4 and dim2 == 4, '2.4.5 - Incorrect answer'
###END HIDDEN TESTS
```

<div class="alert alert-block alert-danger"><b>Additional study material:</b>

* Official Python Documentation - https://docs.python.org/3/tutorial/controlflow.html
* https://realpython.com/python-for-loop/
* Think Python (2nd ed.) - Section 7

<h4>After this Notebook you should be able to:</h4>

- understand the difference between built-in and third-party modules
- use functions from the **`math`** module
- find available functions from any module
- generate an array for the x-axis
- calculate the **`cos`** or **`sin`** of the x-axis
- plot such functions
- understand how to load Python files as modules
- understand conditions with **`if`**, **`elif`** and, **`else`**
- understand the differences between **`list`**, **`tuple`**, and **`dict`**
- slice lists and tuples
- use **`for`** and **`while`** loops
- use the **`break`** and **`continue`** keywords inside of loops
