<figure>
  <IMG SRC="https://raw.githubusercontent.com/mbakker7/exploratory_computing_with_python/master/tudelft_logo.png" WIDTH=250 ALIGN="right">
</figure>

# Modules, conditions, data structures and loops Exercises

## (Searching) Exercise 2.1.1

Write a function, using a function from the <code>math</code> module, to find the greatest common divisor of any two numbers. Your function should return one integer number.

```python
def find_gcd(a, b):
    ...

###BEGIN SOLUTION TEMPLATE=

import math

def find_gcd(a, b):
    gcd = math.gcd(a, b)
    return gcd

###END SOLUTION
print('The greatest common divisor of 2 and 4 is:', find_gcd(2, 4))
```

<!-- grade_id: Q2.1.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert find_gcd(16, 32) == 16, '2.1.1 — Incorrect answer'
###END HIDDEN TESTS
```

## (Searching) Exercise 2.1.2

We can only take and store measurements at a limited number of locations and/or times. But what if we are interested in a value in between, i.e., at a location/time where we do not have a measurement? Then we can use interpolation to estimate that value. A popular, and simple, interpolation technique is <a href="https://en.wikipedia.org/wiki/Linear_interpolation">linear interpolation</a>. 
Your task is to use the module <code>scipy</code> to perform a 1D linear interpolation between a set of known points, where <code>x_known</code> and <code>y_known</code> are arrays with the measured $x$ and $y$ values. Use Google to look up the 1D interpolation function in <code>scipy</code>.<br><br>In the code below you have to replace the <code>...</code> with the correct code.

```python
... # import the module/submodule
def interpolate_inbetween(x_known, y_known, x_predict):
    f = ... # call the 1D interpolation function with proper arguments
    return f(x_predict)

###BEGIN SOLUTION TEMPLATE=
from scipy import interpolate

def interpolate_inbetween(x_known, y_known, x_predict):
    f = interpolate.interp1d(x_known, y_known)
    return f(x_predict)
###END SOLUTION

# Let's try it...
y_predict = interpolate_inbetween([1, 2], [1, 2], 1.5)
print(y_predict)
```

<!-- grade_id: Q2.1.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert abs(interpolate_inbetween([1, 2], [1, 2], 1.5) - 1.5) <= 1e-6, '2.1.2 — Incorrect answer'
###END HIDDEN TESTS
```

## (Searching) Exercise 2.1.3

Now, let's try to measure the running time of a function, for that we will need the <code>time</code> module. Use it to measure the working time of the <code>cool_function()</code> below.

```python
# you do not need to change anything in this cell

def cool_function():
    x = 0
    for i in range(100000000):
        x += 1

```

```python
# complete the code by replacing only the ... to create a function to calculate the running time of any function

... # remember to import the module needed

def measure_time(func):
    t0 = ... # use a function from the time module to store the current time
    func()
    t1 = ... # and again
    return ... # this should calcate the running time

###BEGIN SOLUTION TEMPLATE=
import time

def measure_time(func):
    t0 = time.time()
    func()
    t1 = time.time()
    return t1 - t0
###END SOLUTION

# Let's try it for our cool_function
print('It took {:.2f} seconds for Python to count from 0 to 100000000.'.format(measure_time(cool_function)))
```

<!-- grade_id: Q2.1.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert type(measure_time(cool_function)) == float, '2.1.3 — Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 2.2.1

One of the most crucial applications of <code>if</code> statements is filtering the data from errors and checking whether an error is within a certain limit.<br><br>For example, checking whether the difference between an estimated value and the actual value are within a certain range.<br><br>
Mathematically speaking, this can be expressed as<br><br> $$|\hat{y} - y| < \epsilon$$<br>where $\hat{y}$ is your estimated value, $y$ is the actual value and $\epsilon$ is a certain error threshold.<br><br><br>The function <code>check_error_size()</code> below must do the same — it should return <code>True</code> if the error is within the acceptable range <code>eps</code> and <code>False</code> if it is not.

```python
def check_error_size(estimated_value, true_value, eps):
    ... # your if statement (error in acceptable range)
        return True
    ... # statement if not in acceptable range
        return False

###BEGIN SOLUTION TEMPLATE=
def check_error_size(estimated_value, true_value, eps):
    if abs(estimated_value - true_value) <= eps:
        return True
    else:
        return False
###END SOLUTION

#You can try to check it by yourself by running the function with some self-chosen numbers
check_error(...)
```

<!-- grade_id: Q2.2.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert check_error(0.5, 0.4, 0.2), '2.2.1 - Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 2.2.2

Use the knowledge you have obtained in this Notebook and write your very own function to classify soil samples based on the average grain size. The classification you have to implement is the following:<br>

1.   Clay: avg grain size $<$ 0.002 mm
2.   Silt: 0.002 mm $\leq$ avg grain size $<$ 0.063 mm
3.   Sand: 0.063 mm $\leq$ avg grain size $<$ 2 mm
4.   Gravel: 2 mm $\leq$ avg grain size $<$ 63 mm

Your task is to write the function, which will return the name of the soil type based on the provided average grain size in meters.

```python
def classify_soil(avg_grain_size):
    ...

###BEGIN SOLUTION TEMPLATE=
def classify_soil(avg_grain_size):
    if avg_grain_size < 0.002:
        return 'Clay'
    elif avg_grain_size >= 0.002 and avg_grain_size < 0.063:
        return 'Silt'
    elif avg_grain_size >= 0.063 and avg_grain_size < 2:
        return 'Sand'
    elif avg_grain_size >= 2:
        return 'Gravel'
###END SOLUTION

print(classify_soil(1.5))
```

<!-- grade_id: Q2.2.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert classify_soil(0.0015) == 'Clay' and \
classify_soil(0.05) == 'Silt' and \
classify_soil(0.1) == 'Sand' and \
classify_soil(2.5) == 'Gravel' , '2.2.2 - Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 2.2.3 — Triangle Inequality

Let's imagine that you received a request to manage a web-application, which provides and visualizes <a href="https://en.wikipedia.org/wiki/Interferometric_synthetic-aperture_radar"><b>InSAR</b></a> data of the estimated displacement in an area.<br><br>The app is already quite cool, however, there is always room for improvement. For example, one might want to look into the statistics over the area of a triangle.<br><br>Your goal is to help implement such functionality by checking whether the user selected a valid triangle. You need to do this by checking that the length of <b>each</b> of the sides is smaller than the sum of the other two.

```python
def check_triangle(side_a, side_b, side_c):
    ...

###BEGIN SOLUTION TEMPLATE=
def check_triangle(side_a, side_b, side_c):
    part1 = side_a < (side_b + side_c)
    part2 = side_b < (side_a + side_c)
    part3 = side_c < (side_a + side_b)
    

    if part1 and part2 and part3:
        return "Valid Triangle"
    else:
        return "Invalid Triangle"
###END SOLUTION
```

<!-- grade_id: Q2.2.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert check_triangle(1, 2, 3) == "Invalid Triangle" and check_triangle(3, 4, 5) == "Valid Triangle", '2.2.3 - Incorrect answer'
###END HIDDEN TESTS
```

## (Searching) Exercise 2.2.4

Conditional expression is a way how one can compress an <code>if</code> statement to a more compact (and logical) statement. Your task is to rewrite the below <code>if</code> statement by using the <i>'conditional expression'</i> technique.

```python
y = 0
x = 5

if x % 2 == 0:
    y = x ** 2
else:
    y = x % 3
    

```

```python
# write your code here
y = ...

###BEGIN SOLUTION TEMPLATE=
y = x ** 2 if x % 2 == 0 else x % 3
###END SOLUTION

print(y)
```

## Exercise 2.3.1

Now, let's get down to practice.<br><br>Your first task is to finish the <code>pack_variables()</code> function — a function which will combine all inputs in one list and return it.<br><br>More precisely, this function will receive $5$ arguments as input and you have to return a list with these $5$ elements inside.

```python
def pack_variables(arg1, arg2, arg3, arg4, arg5):
    ...

###BEGIN SOLUTION TEMPLATE=
def pack_variables(arg1, arg2, arg3, arg4, arg5):
    pack = [arg1, arg2, arg3, arg4, arg5]
    return pack
###END SOLUTION

print(pack_variables(1, 2, 4, 22, 7))

```

<!-- grade_id: Q2.3.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert len(pack_variables(221, 22, "2", 22, 7)) == 5 and type(pack_variables(221, 22, "2", 22, 7)) == str, '2.3.1 - Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 2.3.2

Here, you will have to perform a quality assessment on a received list of GNSS measurements. But a simple one.<br><br>You have to check whether the received data has more than $1000$ measurements, in that case we know the receiver was running for a long time without being interrupted.
If it is shorter than $1000$, we assume there were some interruptions. Hence, your function should return a message whether the data is fine or not.

```python
def check_data(measurements):
    ...

###BEGIN SOLUTION TEMPLATE=
def check_data(measurements):
    if len(measurements) < 1000:
        return "Something wrong"
    else:
        return "All fine"
###END SOLUTION

#You can check your code below, where we simulate 1250 measurements
import random 

gnss_data = [random.random() * 2.2e8 for i in range(1250)]
print(check_data(gnss_data))
```

<!-- grade_id: Q2.3.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert check_data([1,2,3]) == "Something wrong" and check_data([0 for i in range(1001)]), '2.1.2 - Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 2.3.3

In this exercises you need to write something opposite to a packager... an unpacker! Sometimes it is easier to work with the most convenient data type, that is best suited for a specific task. In this example, you have to write a function which accepts data stored in a dictionary and saves some data from it. More precisely, you have to select $x$ and $y$ coordinates saved in the <code>input_dict</code> dictionary, and return them as a tuple, with $x$ being the first entry.

```python
# you do not have to change anything in this cell

input_dict = {
    'ID': '334856',
    'operator_name': 'Jarno',
    'observation_amount': '2485',
    'loc_x': [2, 5, 10, 12, -5, 8, 27, 1],
    'loc_y': [6, 1, -5, 15, 4, 8, 0, 10]
}

```

```python
def unpack_dictionary(input_dict):
    ...

###BEGIN SOLUTION TEMPLATE=
def unpack_dictionary(input_dict):
    return (input_dict['loc_x'], input_dict['loc_y'])
###END SOLUTION    

print(unpack_dictionary(input_dict))
```

<!-- grade_id: Q2.3.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert type(unpack_dictionary(input_dict)) == tuple, '2.3.3 - Incorrect answer'
###END HIDDEN TESTS
```

## (Searching) Exercise 2.3.4 — List comprehension

Sometimes you have to create a list of predefined size and fill it with data. This can be done easily using list comprehension.
Write a function using list comprehension, in order to create a list of a predefined size and fill them with zeros.

```python
def create_list(size):
    ...

###BEGIN SOLUTION TEMPLATE=
def create_list(size):
    new_list = [0 for i in range(size)]
    return new_list
###END SOLUTION

print(create_list(5))
```

<!-- grade_id: Q2.3.4, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert len(create_list(5)) == 5 and type(create_list(5)) == list, '2.3.4 - Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 2.4.1

In this exercise you will write your own Celsius to Fahrenheit converter! Your task is to write a function, which will accept the list of temperatures <code>temp_c</code>, in Celsius, and will output a list with the same temperatures, but in Fahrenheit.

```python
# you do not need to change anything in this cell

temp_c = [-1, -1.2, 1.3, 6.4, 11.2, 14.8, 17.8, 17.7, 13.7, 8.5, 4.1, 0.9]
```

```python
def celsius_to_fahrenheit(temp_c):
    ...

###BEGIN SOLUTION TEMPLATE=
def celsius_to_fahrenheit(temp_c):
    temp_f = []

    for i in range(len(temp_c)):
        temp_f.append(temp_c[i] * 9 / 5 + 32) 
    return temp_f
###END SOLUTION

print(celsius_to_fahrenheit(temp_c))
```

<!-- grade_id: Q2.4.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert type(celsius_to_fahrenheit([-1])) == list and abs(celsius_to_fahrenheit([-1])[0] - 30.2) <= 1e-6, '2.4.1 - Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 2.4.2

Your task here is to write a function which will analyze a broadcasting message of the following format: <b>"satellite_ids;date"</b>, where the first part of the message contains unique lowercase letters, each corresponding to a different satellite ID, and the last part contains the date of the message.<br><br>
Here are some examples:
    <b>"agf;06062022"</b>,<b> "abcdefgops;03121999" </b>or<b> "xyz;11112011".</b>
<br><br>Your task is to write a function, which for a provided broadcast message, will count the number of satellites mentioned in the message.


```python
def count_satellites(message):
    ...

###BEGIN SOLUTION TEMPLATE=
def count_satellites(message):
    for i in range(len(message)):
        if message[i] == ';':
            break

    return i
###END SOLUTION

# Check that with the example below you count 5 satellites
print(count_satellites("hallo;12122007"))

```

<!-- grade_id: Q2.4.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert count_satellites("qute;12122007") == 4, '2.4.2 - Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 2.4.3

Here you need to write a function that is able to sort any list consisting only of real numbers, in the descending order. For example, the list $[19, 5, 144, 6]$ becomes $[144, 19, 6, 5]$. Hint: use a built-in function or write your own bubble-sort algorithm to sort the list in ascending order, and then think of a clever way to change the order this list to descending order.

```python
def sort_list(unsorted_list):
    ...

###BEGIN SOLUTION TEMPLATE=
def sort_list(unsorted_list):
    return sorted(unsorted_list)[::-1]
###END SOLUTION


print(sort_list([9, -1, 5, 1, -9, -9]))
```

<!-- grade_id: Q2.4.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert sort_list([3, 1, 2]) == [3, 2, 1], '2.4.3 - Incorrect answer'
###END HIDDEN TESTS
```

Before you continue with exercises 2.4.4 and 2.4.5 we would like to try and facilitate your visualization of how to 'navigate' a 2D object such as a Matrix or a 2D list. All comes down to understanding indexing.

```python
import numpy as np # ignore the use of np.array at this point, it just facilitates the printed
# 'matrix view' below

mat = np.array([['[0][0]','[0][1]','[0][2]','[0][3]'],['[1][0]','[1][1]','[1][2]','[1][3]'],['[2][0]','[2][1]','[2][2]','[2][3]']])
print(mat)
```

Output:
```
[['[0][0]' '[0][1]' '[0][2]' '[0][3]']
 ['[1][0]' '[1][1]' '[1][2]' '[1][3]']
 ['[2][0]' '[2][1]' '[2][2]' '[2][3]']]
```

To facilitate this illustration, the above Matrix has its own index as elements. In a 2D object, the first <b>indexing</b> is the number of the <b>row</b>, the second is the number of the <b>column</b>. Let's break it down.

```python
print(mat[0]) # the 0-th row of the variable 'mat'
```

Output:
```
['[0][0]' '[0][1]' '[0][2]' '[0][3]']
```

```python
print(mat[1]) # the 1-st row of the variable 'mat'; etc..
```

Output:
```
['[1][0]' '[1][1]' '[1][2]' '[1][3]']
```

```python
print(mat[2,3]) # the 2-nd row and 3-th column
```

Output:
```
[2][3]
```

```python
print(mat[:,2]) # all rows and 2-nd column (in other words, only the 2-nd column)
```

Output:
```
['[0][2]' '[1][2]' '[2][2]']
```

```python
print(mat[-1,:]) # last row, all columns (therefore, last row)
```

Output:
```
['[2][0]' '[2][1]' '[2][2]' '[2][3]']
```

```python
# could have also been achieved only by writing mat[-1]
print(mat[-1])
```

Output:
```
['[2][0]' '[2][1]' '[2][2]' '[2][3]']
```

```python
# check the number of rows of the var 'mat'
print(len(mat))
```

Output:
```
3
```

```python
# check the number of columns of the var 'mat'
print(len(mat[0])) # number of elements in row [0]
```

Output:
```
4
```

Note that here we only check the number of elements in the 0-th row, since all rows have same number of elements. However, you may also work with irregular lists, where each 'row' or 'sublist' will have a different number of 'columns'.

Now that you are a bit more familiarized with indexing, let's hop into the last two exercises of this Notebook.

## Exercise 2.4.4

Quite frequently you will have to work with 2D data, which (in Python) is stored, sometimes, in 2 dimensional lists. Your task is to complete a function, which will count elements in an irregular 2D list (irregular means that the amount of elements within each sub-list is different). For that you will have to write a loop inside a loop (or, in other words, a double loop). The code below is almost complete, you just need finish typing the second <code>for</code> loop.

```python
def count_elements(data):
    #initializing counter
    cnt = 0    
    for i in range(len(data)) # looping through the number of rows
        for ... # looping to count elements in each row (or sub-list, in this case)
            cnt += 1
            
    return cnt

###BEGIN SOLUTION TEMPLATE=
def count_elements(data):
    #initializing counter
    cnt = 0    
    for i in range(len(data)):
        for j in range(len(data[i])):
            cnt += 1
            
    return cnt

###END SOLUTION

test = [[1, 2, 3], ["list", 4.5], [9, -1, "another list", 6, 101]]
print(f'Total number of elements inside a test list: {count_elements(test)}')
```

Output:
```
Total amount of elements inside a test list: 10
```

<!-- grade_id: Q2.4.4, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
test_data = [[1], [2, 3], [4, 5, 6], [7, 8, "9", 10]]
assert count_elements(test_data) == 10, '2.4.4 - Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 2.4.5

In this exercise you will perform <a href="https://en.wikipedia.org/wiki/Downsampling_(signal_processing)">downsampling</a> of a provided 'regular' 2D list. Downsampling is a procedure where only a subset of the data is sampled (to reduce its size, for example). Below a visual aid of what downsampling of a 2D list is. Instead of using all the data available, your task is to downsample the 2D list to keep only the data of every $a$-th row and every $b$-th column, including the first element $(a_{0,0})$.<br><br>$$A = \left[\begin{array}{ccccc}
a_{0,0} & \dots & a_{0,b} & \dots & a_{0,2b} & \dots & \dots & a_{0,nb}	& \dots\\
\vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots\\
a_{a,0} & \dots & a_{a,b} & \dots & a_{a,2b} & \dots & \dots & a_{a,nb}	& \dots\\
\vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots\\
a_{2a,0} & \dots & a_{2a,b} & \dots & a_{2a,2b} & \dots & \dots & a_{2a,nb}	& \dots\\
\vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots\\
\vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots\\
a_{ma,0} & \dots & a_{ma,b} & \dots & a_{ma,2b} & \dots & \dots & a_{ma,nb} & \dots \\
\vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots & \vdots
\end{array}\right]$$

<br><i>Hint: Use the index value of each element to realize if they should be included or not in the downsample_data list.<br>Hint2: Remember that $0$ % $X = 0$, for any $X$. (Modulus operator, discussed in Notebook 1).

```python
def downsample(data, a, b):
    downsample_data = []

    # create a for loop over len(data), which is the number of rows
    for ...
        row = []

        # create a for loop over len(data[i]), which is the number of columns
        for ...
            # if an element is to be stored, append row variable (created above)
            ... 
            ...

        # finally append downsample_data      
        if len(row) > 0:
            ...
         
    return ...

print(downsample(insar_data, 3, 4))
###BEGIN SOLUTION TEMPLATE=
def downsample(data, a, b):
    downsample_data = []
    
    for i in range(len(data)):
        row = []
        for j in range(len(data[i])):
            if i % a == 0 and j % b == 0:
                row.append(data[i][j])
        if len(row) > 0:
            downsample_data.append(row)
            
    return downsample_data
###END SOLUTION

import numpy as np

insar_data = [[np.random.rand() for i in range(12)] for i in range(9)]

print(downsample(insar_data, 3, 4))
```

<!-- grade_id: Q2.4.5, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
insar_data = [[np.random.rand() for i in range(16)] for i in range(12)]
downsampled = downsample(insar_data, 3, 4)
dim1 = len(downsampled)
dim2 = len(downsampled[0])
assert dim1 == 4 and dim2 == 4, '2.4.5 - Incorrect answer'
###END HIDDEN TESTS
```