# Pandas Exercises
## Exercise 6.2.1
Finish a function, which with a <code>for</code> loop will get all types saved in series <code>my_series</code>


```python
'''write your code here'''

def list_types(series):
    series_types = "Types inside series:\n"
    for ...
        item_type = ...
        series_types += str(item_type) + '\n'
    
    return series_types
        
###BEGIN SOLUTION TEMPLATE=
def list_types(series):
    series_types = "Types inside series:\n"
    for i in range(len(series)):
        item_type = type(series[i])
        series_types += str(item_type) + '\n'
    
    return series_types
###END SOLUTION

print(list_types(my_series))
```

<!-- grade_id: Q6.2.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
def solution_6_2_1(series):
    series_types = "Types inside series:\n"
    for i in range(len(series)):
        item_type = type(series[i])
        series_types += str(item_type) + '\n'
    
    return series_types

assert list_types(my_series) == solution_6_2_1(my_series), '6.2.1 - Incorrect answer'
###END HIDDEN TESTS
```

<div style="text-align: justify">Note the above exercise could have been solved without a <b><code>for</code></b> loop:</div>

```python
types = my_series.apply(type)
print(types)
```

## (Searching) Exercise 6.3.1

Complete the code below to print the total number of <i>'NaN'</i> values in <code>df3</code>

```python
''' change only the ... below, do not add any other lines to this cell '''

def count_nans(df):
    nan_total = ...
    return nan_total

###BEGIN SOLUTION TEMPLATE=
def count_nans(df):
    nan_total = df.isnull().sum().sum()
    return nan_total
###END SOLUTION

print(df3.head(10))
print(f'\ntotal amount of nans = {count_nans(df3)}')
```

<!-- grade_id: Q6.3.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert count_nans(df3) == 17, '6.3.1 - Incorrect answer'
###END HIDDEN TESTS
```

## (Searching) Exercise 6.3.2

Alter <code>df_632</code> as follows:<br>1) Change the index of the rows to be the column 'name';<br>2) drop the column 'name';<br>3) drop the columns 'test' and 'new_column';<br>4) drop the row named 'test';<br>5) add a new column '% of Earth's crust' with values 'TBD'.<br><br>For your easy reference, your final <code>df_632</code> should look like this:
    
![image.png](01.png)

```python
df_632 = df3.copy()
# write your code here
...

###BEGIN SOLUTION TEMPLATE=

df_632.index = df_632.name
df_632.drop('name',axis=1, inplace=True)
df_632.drop('new_column',axis=1,inplace=True)
df_632.drop('test',inplace=True)
df_632.drop('test',axis=1,inplace=True)
df_632["% of Earth's crust"] = 'TBD'
###END SOLUTION


print(df_632.head(10))
```

<!-- grade_id: Q6.3.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
df_sol = df3.copy()
df_sol.index = df_sol.name
df_sol.drop('name',axis=1, inplace=True)
df_sol.drop('new_column',axis=1,inplace=True)
df_sol.drop('test',inplace=True)
df_sol.drop('test',axis=1,inplace=True)
df_sol["% of Earth's crust"] = 'TBD'

assert df_sol.equals(df_632), '6.3.2 - Incorrect answer.\
Make sure to use inplace=True and axis=1 for removing columns'
###END HIDDEN TESTS
```

## (Searching/Fixing) Exercise 6.3.3

A geologist wrote a line of code to know how many of the listed minerals have a hardness greater than or equal to $3$. However, something is not correct in his code... could you maybe fix it? There is a syntax error as well as a semantic error.

```python
def count_minerals(df, minimal_hardness):
    amount_of_minerals = len(df(df.hardness.gt(minimal_hardness)))
    return amount_of_minerals

###BEGIN SOLUTION TEMPLATE=
def count_minerals(df, minimal_hardness):
    amount_of_minerals = len(df[df.hardness.ge(minimal_hardness)])
    return amount_of_minerals                           
###END SOLUTION            
      
print(df3_cleaned.head(10))
print(f'Amount of minerals with hardness >= 3: {count_minerals(df3_cleaned, 3)}')
```

<!-- grade_id: Q6.3.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert count_minerals(df3_cleaned, 4) == 4, '6.3.3 - Incorrect answer'
###END HIDDEN TESTS
```

The above exercise might look like a silly one, but take some time to analyze it carefully... what is happening there? Break down each part of that line of code to understand it properly.

## (Searching) Exercise 6.5.1

A geologist is interested in the tallest mountain chains around Earth. For that, he created the <code>tallest_mountains.csv</code> table, containing information on all mountains above $8000$ meters. Your task is to do the following $5$ assignments.<br><br>1) Read the <code>tallest_mountains.csv</code> file.<br>2) What are the names of the columns?<br>3) What is the height of the tallest mountain?<br>4) What is the row number (or index) of this mountain?<br>5) What is the name of the tallest mountain in the dataset?<br><br><b>Hard-coded answers won't be accepted. Use <code>pandas</code> functions to answer the questions.</b><br><br><i>Source for this and next exercises: <a href="https://en.wikipedia.org/wiki/List_of_mountains_by_elevation">Wikipedia

```python
''' write your code here, do not change any variable name '''

df8000 = ... # 1) read the file
cols = ... # 2) name of columns
max_height = ... # 3) height of tallest mountain
indxmax = ... # 4) index of the row containing 'max_height'
tallest_mountain = ... # 5) name of the tallest mountain

###BEGIN SOLUTION TEMPLATE=
df8000 = pd.read_csv(path + 'tallest_mountains.csv')
cols = df8000.columns
max_height = df8000['Metres'].max() 
indxmax = df8000['Metres'].idxmax() 
tallest_mountain = df8000.loc[indxmax,'Mountain']
###END SOLUTION

print(df8000.head())
```

<!-- grade_id: Q6.5.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
df_sol = pd.read_csv(path + 'tallest_mountains.csv')
cols_sol = df_sol.columns
max_height_sol = df_sol['Metres'].max() 
indxmax_sol = df_sol['Metres'].idxmax() 
tallest_mountain_sol = df_sol.loc[indxmax_sol,'Mountain']

assert df_sol.equals(df8000) and cols_sol.equals(cols) and max_height_sol == max_height and \
        indxmax_sol == indxmax and tallest_mountain_sol == tallest_mountain, \
    '6.5.1 - Incorrect answer, did you use idxmax for the 4th problem?'
###END HIDDEN TESTS
```

## (Searching) Exercise 6.5.2

Now, our geologist friend got another table with all mountains above $7000$ meters, it's in the <code>mountains_above_7000m_copy.csv</code> file. Your task is to do the following $6$ assignments.<br><br>1) Read the <code>mountains_above_7000m_copy.csv</code> file.<br>2) Remove the column describing the mountain range they belong.<br>3) Append <code>df8000</code> to this.<br>4) Fix the row-indexing issue.<br>5) Create a <code>Series</code> indicating which entries of 'Feet' are missing.<br><i>Hint: Use the <code>.isnull()</code> function</i><br>6) Add the values that are missing in the 'Feet' column, use $1$ mtr = $3.28084$ feet.<br><i>Hint2: Use the <code>.mask()</code> function, more info <a href="https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.mask.html">here</a>, or the <code>.fillna()</code> function.

```python
''' write your code here, do not change any variable name '''

df7000 = ... # 1
df7000_norange = ... # 2
df_concat = ... # 3

#check what you have now
print(df_concat)
```

```python
''' write the rest of your code here, do not change any variable name '''

df_reset = ... # 4
missing_feet_series = ... # 5
df_reset["Feet"] = ... # 6
print(df_reset) 

###BEGIN SOLUTION TEMPLATE=
df7000 = pd.read_csv(path + 'mountains_above_7000m_copy.csv', encoding_errors='ignore') # 1
df7000_norange = df7000.drop('Range', axis=1) # 2
df_concat = pd.concat([df7000_norange,df8000]) # 3
df_reset = df_concat.reset_index(drop=True) # 4
missing_feet_series = df_reset["Feet"].isnull() # 5
df_reset["Feet"] = df_reset["Feet"].mask(missing_feet_series, df_reset["Metres"]*3.28084) # 6
###END SOLUTION
```

<!-- grade_id: Q6.5.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
df_sol = pd.read_csv(path + 'tallest_mountains.csv')

df7000_sol = pd.read_csv(path + 'mountains_above_7000m_copy.csv', encoding_errors='ignore') # 1
df7000_norange_sol = df7000_sol.drop('Range', axis=1) # 2
df_concat_sol = pd.concat([df7000_norange_sol,df_sol]) # 3
df_reset_sol = df_concat_sol.reset_index(drop=True) # 4
missing_feet_series_sol = df_reset_sol["Feet"].isnull() # 5
with_feet_series_sol = df_reset_sol["Feet"].mask(missing_feet_series_sol, df_reset_sol["Metres"]*3.28084) # 6
df_reset_sol["Feet"] = with_feet_series_sol

or_df7000_sol = pd.read_csv(path + 'mountains_above_7000m_copy.csv', encoding_errors='ignore') # 1
or_df7000_norange_sol = or_df7000_sol.drop('Range', axis=1) # 2
or_df_concat_sol = pd.concat([df_sol,or_df7000_norange_sol]) # 3
or_df_reset_sol = or_df_concat_sol.reset_index(drop=True) # 4
or_missing_feet_series_sol = or_df_reset_sol["Feet"].isnull() # 5
or_with_feet_series_sol = or_df_reset_sol["Feet"].mask(or_missing_feet_series_sol, or_df_reset_sol["Metres"]*3.28084) # 6
or_df_reset_sol["Feet"] = or_with_feet_series_sol

assert df_reset_sol.equals(df_reset) or or_df_reset_sol.equals(df_reset), '6.5.2 - Incorrect answer'
###END HIDDEN TESTS
```

## (Searching) Exercise 6.5.3

Now that our geologist friend has all this information, he wants to know how many of these mountains are claimed by China. Can you help him out? <i><br><br>Hint: you need to find out which elements in the appropriate column contains the string 'China'.

```python
''' write your code here, do not change any variable name '''

# you only need one line of code
china_mountains = ...

###BEGIN SOLUTION TEMPLATE=
china_mountains = df_reset["Location and Notes"].str.contains("China", case=True).sum()
###END SOLUTION

print(china_mountains)
```

<!-- grade_id: Q6.5.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
df_sol = pd.read_csv(path + 'tallest_mountains.csv')

df7000_sol = pd.read_csv(path + 'mountains_above_7000m_copy.csv', encoding_errors='ignore') # 1
df7000_norange_sol = df7000_sol.drop('Range', axis=1) # 2
df_concat_sol = pd.concat([df7000_norange_sol,df_sol]) # 3
df_reset_sol = df_concat_sol.reset_index(drop=True) # 4
missing_feet_series_sol = df_reset_sol["Feet"].isnull() # 5
with_feet_series_sol = df_reset_sol["Feet"].mask(missing_feet_series_sol, df_reset_sol["Metres"]*3.28084) # 6
df_reset_sol["Feet"] = with_feet_series_sol

china_mountains_sol = df_reset_sol["Location and Notes"].str.contains("China", case=True).sum()
assert china_mountains_sol == china_mountains, '6.5.3 - Incorrect answer'
###END HIDDEN TESTS
```