# Pandas Exercises
## Exercise 6.2.1
Finish a function, which with a <code>for</code> loop will get all types saved in series <code>my_series</code>

_Type your code where the three (*`...`*) dots are placed. Do not change the name of the variables._

```python
import pandas as pd

my_list = ['begin', 2, 3/4, "end"]
my_series = pd.Series(data=my_list)

def list_types(series):
    series_types = "Types inside series:\n"
    for ...
        item_type = ...
        series_types += str(item_type) + '\n'
    
    return series_types
        
print(list_types(my_series))
```

Output:
```
Cell In[7], line 8
    for ...
           ^
SyntaxError: invalid syntax
```

````{admonition} check your answer!
Run the cell bellow to check if you got the correct answer
````

```python
notebook_6_check_answer(0)
```

Output:
```
Correct answer :D
```

## (Searching) Exercise 6.3.1

Complete the code below to print the total number of <i>'NaN'</i> values in <code>mineral_properties</code>

_Change only the (*`...`*) below, do not add any other lines to this cell_
 

```python
import numpy as np
import pandas as pd

file_location = ("https://raw.githubusercontent.com/TUDelft-CITG/"
                "learn-python/mike/book/06/Theory/")
mineral_properties = pd.read_csv(file_location + 'mineral_properties.txt', 
                  skiprows=1,skipinitialspace=True)
mineral_properties['new_column'] = np.nan

def count_nans(df):
    nan_total = ...
    return nan_total

print(mineral_properties.head(3))
print(f'\ntotal amount of nans = {count_nans(mineral_properties)}')
```

Output:
```
name  hardness  sp. gr. cleavage  new_column
0  Amphibole      5.50     2.80      Two         NaN
1    Biotite      2.75     3.00      One         NaN
2    Calcite      3.00     2.72    Three         NaN

total amount of nans = Ellipsis
```

````{admonition} check your answer!
Run the cell bellow to check if you got the correct answer
````

```python
notebook_6_check_answer(1)
```

```python
import pandas as pd

def solution_6_2_1(series):
    series_types = "Types inside series:\n"
    for i in range(len(series)):
        item_type = type(series[i])
        series_types += str(item_type) + '\n'
    return series_types

def notebook_6_check_answer(question_number):

    if question_number == 0:
        my_list = ['begin', 2, 3/4, "end"]
        my_series = pd.Series(data=my_list)
        try: 
            assert list_types(
                my_series) == solution_6_2_1(
                my_series), '6.2.1 - Incorrect answer'
            print('Correct answer :D')
        except AssertionError:
            print('Incorrect answer :(')

    elif question_number == 1:
        try: 
            assert count_nans(
                mineral_properties) == 12,'Incorrect'
            print('Correct answer :D')
        except AssertionError:
            print('Incorrect answer :(')

    elif question_number == 2:
        try: 
            assert count_minerals(
                mineral_properties, 4) == 7, 'Incorrect'
            print('Correct answer :D')
        except AssertionError:
            print('Incorrect answer :(')
    
    elif question_number == 3:
        try: 
            assert get_abbreviation(   
            ) == "AES", '3.4.1 - Incorrect answer'
            print('Correct answer :D')
        except AssertionError:
            print('Incorrect answer :(')

    elif question_number == 4:
        try: 
            assert "B[3] = 8" in create_string_from_lists(
            ), '3.4.2 - Incorrect answer'
            print('Correct answer :D')
        except AssertionError:
            print('Incorrect answer :(')

    elif question_number == 5:
        try: 
            assert factorial(
                5) == 120, '3.4.3 - Incorrect answer'
            print('Correct answer :D')
        except AssertionError:
            print('Incorrect answer :(')
```

## (Searching/Fixing) Exercise 6.3.3

A geologist wrote a line of code to know how many of the listed minerals have a hardness greater than or equal to $3$. However, something is not correct in his code... could you maybe fix it? There is a syntax error as well as a semantic error.

```python
def count_minerals(df, minimal_hardness):
    amount_of_minerals = len(df(df.hardness.gt(minimal_hardness)))
    return amount_of_minerals
             
print(mineral_properties.head(3))
print(f'Amount of minerals with hardness >= 3: {count_minerals(mineral_properties, 3)}')
```

<!-- grade_id: Q6.3.3, points: 1, grade: true, solution: false -->
```python
notebook_6_check_answer(2)
```

The above exercise might look like a silly one, but take some time to analyze it carefully... what is happening there? Break down each part of that line of code to understand it properly.

## (Searching) Exercise 6.5.1

A geologist is interested in the tallest mountain chains around Earth. For that, he created the <code>tallest_mountains.csv</code> table, containing information on all mountains above $8000$ meters. Your task is to do the following $5$ assignments.
1. Read the <code>tallest_mountains.csv</code> file.
2. What are the names of the columns?
3. What is the height of the tallest mountain?
4. What is the row number (or index) of this mountain?
5. What is the name of the tallest mountain in the dataset?

Use <code>pandas</code> functions to answer the questions.

Source for this and next exercises: <a href="https://en.wikipedia.org/wiki/List_of_mountains_by_elevation">Wikipedia

Write your code here, do not change any variable name


```python
file_location = ("https://raw.githubusercontent.com/TUDelft-CITG/"
                "learn-python/mike/book/06/Exercises/")

mountains_8000 = ... # 1) read the csv file
cols = ... # 2) name of columns
max_height = ... # 3) height of tallest mountain
index_max = ... # 4) index of the row containing 'max_height'
tallest_mountain = ... # 5) name of the tallest mountain

###BEGIN SOLUTION TEMPLATE=
mountains_8000 = pd.read_csv(file_location + 'tallest_mountains.csv')
cols = mountains_8000.columns
max_height = mountains_8000['Metres'].max() 
index_max = mountains_8000['Metres'].idxmax() 
tallest_mountain = mountains_8000.loc[index_max,'Mountain']
###END SOLUTION

print(mountains_8000.head())
```

<!-- grade_id: Q6.5.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
df_sol = pd.read_csv(file_location + 'tallest_mountains.csv')
cols_sol = df_sol.columns
max_height_sol = df_sol['Metres'].max() 
indxmax_sol = df_sol['Metres'].idxmax() 
tallest_mountain_sol = df_sol.loc[indxmax_sol,'Mountain']

assert df_sol.equals(mountains_8000) and cols_sol.equals(cols) and max_height_sol == max_height and \
        indxmax_sol == index_max and tallest_mountain_sol == tallest_mountain, \
    '6.5.1 - Incorrect answer, did you use idxmax for the 4th problem?'
###END HIDDEN TESTS
```

## (Searching) Exercise 6.5.2

Now, our geologist friend got another table with all mountains above $7000$ meters, it's in the <code>mountains_above_7000m.csv</code> file. Your task is to do the following $6$ assignments.
1. Read the <code>mountains_above_7000m.csv</code> file.
2. Remove the column describing the mountain range they belong.
3. Append <code>mountains_8000</code> to this.
4. Fix the row-indexing issue.
5. Create a <code>Series</code> indicating which entries of 'Feet' are missing.<br><i>Hint: Use the <code>.isnull()</code> function</i>
6. Add the values that are missing in the 'Feet' column, use $1$ mtr = $3.28084$ feet.<br><i>Hint2: Use the <code>.mask()</code> function, more info <a href="https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.mask.html">here</a>, or the <code>.fillna()</code> function.

Write your code for tasks 1,2 and 3 in teh cell below, do not change any variable name

```python
file_location = ("https://raw.githubusercontent.com/TUDelft-CITG/"
                "learn-python/mike/book/06/Exercises/")

mountains_7000 = ... # 1
mountains_7000_norange = ... # 2
df_concat = ... # 3

print(df_concat)
```

Write your code for the remaining tasks in teh cell below, do not change any variable name

```python

df_reset = ... # 4
missing_feet_series = ... # 5
df_reset["Feet"] = ... # 6
print(df_reset) 

###BEGIN SOLUTION TEMPLATE=
df7000 = pd.read_csv(file_location + 'mountains_above_7000m_copy.csv', encoding_errors='ignore') # 1
df7000_norange = df7000.drop('Range', axis=1) # 2
df_concat = pd.concat([df7000_norange,df8000]) # 3
df_reset = df_concat.reset_index(drop=True) # 4
missing_feet_series = df_reset["Feet"].isnull() # 5
df_reset["Feet"] = df_reset["Feet"].mask(missing_feet_series, df_reset["Metres"]*3.28084) # 6
###END SOLUTION
```

<!-- grade_id: Q6.5.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
df_sol = pd.read_csv(file_location + 'tallest_mountains.csv')

df7000_sol = pd.read_csv(file_location + 'mountains_above_7000m_copy.csv', encoding_errors='ignore') # 1
df7000_norange_sol = df7000_sol.drop('Range', axis=1) # 2
df_concat_sol = pd.concat([df7000_norange_sol,df_sol]) # 3
df_reset_sol = df_concat_sol.reset_index(drop=True) # 4
missing_feet_series_sol = df_reset_sol["Feet"].isnull() # 5
with_feet_series_sol = df_reset_sol["Feet"].mask(missing_feet_series_sol, df_reset_sol["Metres"]*3.28084) # 6
df_reset_sol["Feet"] = with_feet_series_sol

or_df7000_sol = pd.read_csv(file_location + 'mountains_above_7000m_copy.csv', encoding_errors='ignore') # 1
or_df7000_norange_sol = or_df7000_sol.drop('Range', axis=1) # 2
or_df_concat_sol = pd.concat([df_sol,or_df7000_norange_sol]) # 3
or_df_reset_sol = or_df_concat_sol.reset_index(drop=True) # 4
or_missing_feet_series_sol = or_df_reset_sol["Feet"].isnull() # 5
or_with_feet_series_sol = or_df_reset_sol["Feet"].mask(or_missing_feet_series_sol, or_df_reset_sol["Metres"]*3.28084) # 6
or_df_reset_sol["Feet"] = or_with_feet_series_sol

assert df_reset_sol.equals(df_reset) or or_df_reset_sol.equals(df_reset), '6.5.2 - Incorrect answer'
###END HIDDEN TESTS
```

## (Searching) Exercise 6.5.3

Now that our geologist friend has all this information, he wants to know how many of these mountains are claimed by China. Can you help him out? <i><br><br>Hint: you need to find out which elements in the appropriate column contains the string 'China'.

```python
''' write your code here, do not change any variable name '''

# you only need one line of code
china_mountains = ...

###BEGIN SOLUTION TEMPLATE=
china_mountains = df_reset["Location and Notes"].str.contains("China", case=True).sum()
###END SOLUTION

print(china_mountains)
```

<!-- grade_id: Q6.5.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
df_sol = pd.read_csv(file_location + 'tallest_mountains.csv')

df7000_sol = pd.read_csv(file_location + 'mountains_above_7000m_copy.csv', encoding_errors='ignore') # 1
df7000_norange_sol = df7000_sol.drop('Range', axis=1) # 2
df_concat_sol = pd.concat([df7000_norange_sol,df_sol]) # 3
df_reset_sol = df_concat_sol.reset_index(drop=True) # 4
missing_feet_series_sol = df_reset_sol["Feet"].isnull() # 5
with_feet_series_sol = df_reset_sol["Feet"].mask(missing_feet_series_sol, df_reset_sol["Metres"]*3.28084) # 6
df_reset_sol["Feet"] = with_feet_series_sol

china_mountains_sol = df_reset_sol["Location and Notes"].str.contains("China", case=True).sum()
assert china_mountains_sol == china_mountains, '6.5.3 - Incorrect answer'
###END HIDDEN TESTS
```