<figure>
  <IMG SRC="https://raw.githubusercontent.com/mbakker7/exploratory_computing_with_python/master/tudelft_logo.png" WIDTH=250 ALIGN="right">
</figure>

# Objects and References Exercises

## (Fixing) Exercise 4.1.1

Let's assume that you develop a login interface for an international oil company; which, as the title international suggests, works internationally. Thus, a lot of clients worldwide expect to see data in their favorite units. The company has hired an intern to output temperatures in Celsius and Fahrenheit, from the provided temperature in Kelvin. But something went wrong, the code doesn't work and now the intern is gone. So, the burden of fixing the code is, yet again, put on your shoulders.

```python
import time

def get_display_temperature(temp_k):
    # copying temporarily temp_k to temp_c
    temp_c = temp_k
    
    # converting kelvin to celsius
    for i in range(len(temp_c)):
        temp_c[i] = temp_c[i] - 273.15
        
    # copying temporarily temp_k to temp_f
    temp_f = temp_k
    
    # converting kelvin to fahrenheit
    for i in range(len(temp_f)):
        temp_f[i] = (temp_f[i] - 273.15) * (9 / 5) + 32
    
    # now, creating display messages from the converted temperatures
    display_messages = []
    for i in range(len(temp_k)):
        msg = f"{temp_c[i]:<10.3f}°C | {temp_f[i]:<10.3f}°F (ID={i})"
        display_messages.append(msg)
        
    return display_messages

###BEGIN SOLUTION TEMPLATE=
def get_display_temperature(temp_k):
    #copying temporarily temp_k to temp_c
    temp_c = temp_k.copy()
    
    #converting kelvins to celsius
    for i in range(len(temp_c)):
        temp_c[i] = temp_c[i] - 273.15
        
    #copying temporarily temp_k to temp_f
    temp_f = temp_k.copy()
    
    #converting kelvins to farenheits
    for i in range(len(temp_f)):
        temp_f[i] = (temp_f[i] - 273.15) * (9 / 5) + 32
    
    #now, creating display messages from the converted temperatures
    display_messages = []
    for i in range(len(temp_k)):
        msg = f"{temp_c[i]:<10.3f}°C | {temp_f[i]:<10.3f}°F (ID={i})"
        display_messages.append(msg)
        
    return display_messages
###END SOLUTION

def update_screen_text(messages):
    for msg in messages:
        print(msg, end='\r')
        time.sleep(1)
        

temp_k = [300.67, 277.56, 315.88, 307.87, 100]
messages = get_display_temperature(temp_k)
update_screen_text(messages)
```

<!-- grade_id: Q4.1.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert get_display_temperature([100]) == ["-173.150  °C | -279.670  °F (ID=0)"], '4.1.1 - Incorrect answer'
###END HIDDEN TESTS
```

## (Fixing) Exercise 4.1.2

The main philosophy of programming is avoiding redundancy — you shouldn't write the same batch code two or more times if you can just create a function out of it. The same can be applied to many other things — don't start a new assignment from scratch, if you can re-use a formatting template from your previous assignments. In the cell with code below, you see an example of preparing a template for observation of some satellite, which can work in 2 different modes, and thus will have a slightly different template for each of the modes. The code below is supposed to do that and it was even tested!

```python
import time

def prepare_template(default_bands, observation_mode):  
    #creating metadata for the upcoming observations
    template = {'time': time.ctime(time.time()),
               'observation_mode': observation_mode,
               'bands': default_bands}
    
    #adding additional bands for the extended mode
    if observation_mode == 'normal':
        #no need to add bands
        pass
    elif observation_mode == 'extended':
        template['bands'] += ['B8', 'B8A']
    else:
        #if the mode is unknonw - raise a RuntimeError
        raise RuntimeError(f'Failed to identify observation mode: {observation_mode}')
        
    return template

###BEGIN SOLUTION TEMPLATE=
def prepare_template(default_bands, observation_mode):  
    #creating metadata for the upcoming observations
    template = {'time': time.ctime(time.time()),
               'observation_mode': observation_mode,
               'bands': default_bands.copy()}
    
    #adding additional bands for the extended mode
    if observation_mode == 'normal':
        #no need to add bands
        pass
    elif observation_mode == 'extended':
        template['bands'] += ['B8', 'B8A']
    else:
        #if the mode is unknonw - raise a RuntimeError
        raise RuntimeError(f'Failed to identify observation mode: {observation_mode}')
        
    return template
###END SOLUTION

def print_dict(my_dict):
    
    for key, value in my_dict.items():
        print(f'{key} -> {value}')  
    #print('-'*10)

"set of default bands, don't change!!"
default_bands = ['B1', 'B2', 'B3', 'B4', 'B5', 'B6', 'B7']

#testing the function
#test 1. - normal mode
#Expected behaviour: a dictionary with 3 items and only default bands are inside
print('-------------------')
print('test 1')
temp_norm = prepare_template(default_bands, 'normal')
print_dict(temp_norm)
print('-------------------')
print('test 2')
#test 2. - extended mode
#Expected behaviour: a dictionary with 3 items and extended list of bands
temp_ext = prepare_template(default_bands, 'extended')
print_dict(temp_ext)
print('-------------------')
print('test 3')
#test 3. - any other observation mode
#Expected behaviour - Runtime error
temp_error = prepare_template(default_bands, 'should raise an error, right?')
```

Output:
```
-------------------
test 1
time -> Tue Jul 12 07:48:28 2022
observation_mode -> normal
bands -> ['B1', 'B2', 'B3', 'B4', 'B5', 'B6', 'B7']
-------------------
test 2
time -> Tue Jul 12 07:48:28 2022
observation_mode -> extended
bands -> ['B1', 'B2', 'B3', 'B4', 'B5', 'B6', 'B7', 'B8', 'B8A']
-------------------
test 3
<ipython-input-6-84a3fb455652> in <module>
     66 #test 3. - any other observation mode
     67 #Expected behaviour - Runtime error
---> 68 temp_error = prepare_template(default_bands, 'should raise an error, right?')

<ipython-input-6-84a3fb455652> in prepare_template(default_bands, observation_mode)
     35     else:
     36         #if the mode is unknonw - raise a RuntimeError
---> 37         raise RuntimeError(f'Failed to identify observation mode: {observation_mode}')
     38 
     39     return template

RuntimeError: Failed to identify observation mode: should raise an error, right?
```

Looks good, however, if you rearrange the extended and normal tests, it won't work anymore, look at the output below to see this.

```python
#set of default bands, don't change!!!
default_bands = ['B1', 'B2', 'B3', 'B4', 'B5', 'B6', 'B7']

#testing the function
#test 1. - extended mode
#Expected behaviour: a dictionary with 3 items and extended list of bands
temp_ext = prepare_template(default_bands, 'extended')
print('-------------------')
print('test 1')
print_dict(temp_ext)
print('-------------------')
print('test 2')
#test 2. - normal mode
#Expected behaviour: a dictionary with 3 items and only default bands are inside
temp_norm = prepare_template(default_bands, 'normal')
print_dict(temp_norm)
print('-------------------')
```

Output:
```
-------------------
test 1
time -> Tue Jul 12 07:49:01 2022
observation_mode -> extended
bands -> ['B1', 'B2', 'B3', 'B4', 'B5', 'B6', 'B7', 'B8', 'B8A']
-------------------
test 2
time -> Tue Jul 12 07:49:01 2022
observation_mode -> normal
bands -> ['B1', 'B2', 'B3', 'B4', 'B5', 'B6', 'B7']
-------------------
```

Obviously, the order of tests shouldn't matter, so, most likely, the problem is with the function itself. Please fix the function, in the cell above, so that the second set of tests also works (you can comment out the first set of tests if you want).

<!-- grade_id: Q4.1.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
default_bands = ['B1', 'B2', 'B3', 'B4', 'B5', 'B6', 'B7']
d1 = prepare_template(default_bands, 'normal')
d2 = prepare_template(default_bands, 'extended')
d3 = prepare_template(default_bands, 'normal')

assert d2['bands'] == ['B1', 'B2', 'B3', 'B4', 'B5', 'B6', 'B7', 'B8', 'B8A'] and \
        d3['bands'] == ['B1', 'B2', 'B3', 'B4', 'B5', 'B6', 'B7'], '4.1.2 - Incorrect answer'
###END HIDDEN TESTS
```