<figure>
  <IMG SRC="https://raw.githubusercontent.com/mbakker7/exploratory_computing_with_python/master/tudelft_logo.png" WIDTH=250 ALIGN="right">
</figure>


# Advanced strings and functions. Files and debugging Exercises.

## Exercise 3.1.1

Use the <code>print()</code> function and an f-string to make a statement about some car using all variables stored in the <code>car_info</code> dictionary.

You do not have to change anything in this cell:

Type your code here where the three (*`...`*) dots are placed. Do not change the variable name nor changes in variable `car_info`.


```python
car_info = {
    'top_speed' : '229', #km/h
    'type' : 'Opel Astra'
}

message = ...

```

<!-- grade_id: Q3.1.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
try: 
    assert car_info['top_speed'] in message and car_info['type'] in message, 'Incorrect answer.'
    print('Correct answer :)')
except AssertionError:
    print('Incorrect answer :(')
###END HIDDEN TESTS
```

Output:
```
Incorrect answer :(
```

## Exercise 3.1.2

Print the value of <b>π</b> using $4$ decimal places.

```python
# write your code here
rounded_pi = ...

###BEGIN SOLUTION TEMPLATE=
rounded_pi = f"{pi:.4f}"
###END SOLUTION
print(rounded_pi)
```

<!-- grade_id: Q3.1.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert rounded_pi == "3.1416", '3.1.2 - Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 3.2.1

Write a lambda function that converts a number from degrees to radians.

```python
DegToRad = ...

###BEGIN SOLUTION TEMPLATE=
DegToRad = lambda theta : theta * (pi / 180)
###END SOLUTION

Angle = 180 # Degrees
print(f"An angle of {Angle} Degrees is equal to {DegToRad(180):.3f} radians")
```

<!-- grade_id: Q3.2.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert type(DegToRad) == type(lambda x: x) and abs(DegToRad(90) - pi/2) <= 1e-6, '3.2.1 - Incorrect answer'
###END HIDDEN TESTS
```

## Exercise 3.2.2

Write a lambda function that takes four inputs: $(x_1, y_1, x_2, y_2)$ and computes the <a href="https://en.wikipedia.org/wiki/Euclidean_distance">Euclidian distance</a> between point 1 $(x_1,y_1)$ and point 2 $(x_2,y_2)$.


```python
Distance = ...

###BEGIN SOLUTION TEMPLATE=
Distance = lambda x1, y1, x2, y2: sqrt((x1 - x2) ** 2 + (y1 - y2) ** 2)
###END SOLUTION

x1, y1 = 1, 1
x2, y2 = 2, 2

print(f"Distance between points ({x1}, {y1}) and ({x2}, {y2}) is {Distance(x1, y1, x2, y2):.3f}")
```

<!-- grade_id: Q3.2.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert abs(Distance(1, 1, 3, 3) - 2 * sqrt(2)) <= 1e-6, '3.2.2 - Incorrect answer'
###END HIDDEN TESTS
```

## (Fixing) Exercise 3.3.1

An engineer from Shell wrote a piece of code to find the file with updated logs. The file has a <code>*.txt</code> extension. It is located in the Vocareum environment and can be accessed using a relative or absolute path. The engineer wrote a piece of code to read the file and print its contents but something went wrong.


```python
FilePath = "./resource/lib/publicdata/UpdatedLogs.txt"# Path of the file + (filename.txt)

###BEGIN SOLUTION TEMPLATE=
FilePath = "./resource/lib/publicdata/updated_logs.txt"
###END SOLUTION


txt_File = open(FilePath, mode="r")                         # Open the file in read (r) mode
File_Contents = txt_File.read()                             # Read the contents of the file
txt_File.close()                                            # Close the file (good practice to do so!)

print(File_Contents)                                        # Print the file
```

<!-- grade_id: Q3.3.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert FilePath == "./resource/lib/publicdata/updated_logs.txt", '3.3.1 - Incorrect answer'
###END HIDDEN TESTS
```

## (Fixing) Exercise 3.4.1

Fix the syntax errors so it prints "AES" without removing the variable that holds it. You'll need to fix 2 errors.

```python
def get_abbreviation():
    my abbreviation = "AES"
     return my_abbreviation
        
###BEGIN SOLUTION TEMPLATE=
def get_abbreviation():
    my_abbreviation = "AES"
    return my_abbreviation
###END SOLUTION
        
print(get_abbreviation())
```

<!-- grade_id: Q3.4.1, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert get_abbreviation() == "AES", '3.4.1 - Incorrect answer'
###END HIDDEN TESTS
```

## (Fixing) Exercise 3.4.2

Solve the runtime error so all values of $B$ are printed. The output after running <code>create_string_from_lists()</code> should be the following:<br>
<code>
A[0] = 2
B[0] = 5
A[1] = 3
B[1] = 6
A[2] = 4
B[2] = 7
B[3] = 8
</code>

```python
def create_string_from_lists():
    s = ""
    
    A = [2, 3, 4]
    B = [5, 6, 7, 8]
    for i in range(4):
        s += f"A[{i}] = {A[i]}\n"
        s += f"B[{i}] = {B[i]}\n"
    print(s)

###BEGIN SOLUTION TEMPLATE=
def create_string_from_lists():
    s = ""
    
    A = [2, 3, 4]
    B = [5, 6, 7, 8]
    for i in range(4):
        
        if i < 3:
            s += f"A[{i}] = {A[i]}\n"
            s += f"B[{i}] = {B[i]}\n"
        else:
            s += f"B[{i}] = {B[i]}\n"
    print(s)   
###END SOLUTION

create_string_from_lists()
```

<!-- grade_id: Q3.4.2, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert "B[3] = 8" in create_string_from_lists(), '3.4.2 - Incorrect answer'
###END HIDDEN TESTS
```

## (Fixing) Exercise 3.4.3

Find the semantic error in this function:

```python
def factorial(x):
    "returns the factorial of x"
    if x == 0:
        return 1
    else: 
        return x ** factorial(x-1)

###BEGIN SOLUTION TEMPLATE=
def factorial(x):
    "returns the factorial of x"
    if x == 0:
        return 1
    else: 
        return x * factorial(x-1)
###END SOLUTION

factorial(4)
```

<!-- grade_id: Q3.4.3, points: 1, grade: true, solution: false -->
```python
###BEGIN HIDDEN TESTS
assert factorial(5) == 120, '3.4.3 - Incorrect answer'
###END HIDDEN TESTS
```